Ubiquitin ligase and signalling hub MYCBP2 is required for efficient EPHB2 tyrosine kinase receptor function

Author Response

The following is the authors’ response to the original reviews.

We thank the reviewers for their thoughtful assessment of our work and their valuable critiques which we will address in the “Recommendations for the authors” section below. In particular, we appreciate Reviewer #3 noting the value of the C. elegans model system and our efforts to bridge models with our study. We agree with the reviewer that there is a need to clarify the rationale, presentation and interpretation of our results. We have substantially revised the text in our manuscript and Figure legend to address this issue, and provided extensive new commentary and citations to lay out the logic behind our experiments. Indeed, it was our oversight not being more thorough about this initially. We have further adjusted our conclusions to be less unequivocal. Finally, we added an RPM-1 signaling diagram (Fig. 8A) to more clearly annotate the players in the RPM-1/MYCBP2 signaling network that were evaluated genetically in Fig. 8. Importantly, we provide clearer commentary on how genetic enhancer effects with known RPM-1 binding proteins and the absence of genetic suppression in vab-1/Eph receptor double mutants with components of the RPM-1/FSN-1 ubiquitin ligase complex are consistent with the biochemical finding that MYCBP2 stabilizes but does not degrade EphB2. Text edits reflecting these points are in the abstract, the C. elegans results section starting on line 411, and the discussion on lines 499, 502-504 and 541.

Following extensive discussions between the three reviewers, all three agree that the C. elegans data, as presented, does not add to, and in fact might harm, your bottom line. Our combined suggestion is to take this data out unless you plan to improve it substantially. All reviewers are perplexed by Figure 2F and the presumed interactions of cytosolic proteins with the extracellular domain of EPHB2. At the very least, please provide some suggestions/model/interpretation.

We have adjusted our manuscript substantially to address this. Please see detailed comments in the individual Reviewer sections below.

We would like to thank the reviewers for their thorough examination of our manuscript, constructive criticisms, and helpful suggestions.

Reviewer #1 (Recommendations For The Authors):

The work is extensive in my view, and mostly of high quality. See minor comments on some of the figures below.

Thank you very much.

Two more major comments :

I don't think the C. elegans work adds to - in fact I think it hurts - the statement that this regulatory mechanism is specific to EphB2. I would advise the authors to take it out.

We agree that C. elegans has a sole Eph receptor called VAB-1 and is therefore not a specific model for EPH2B. However, testing MYCBP2 specificity for EPHB2 was not the goal or our perceived value for the C. elegans experiments. We now clarify this in the text of the Results section.

Rather, we are providing evidence that the C. elegans ephrin receptor interacts genetically with known MYCBP2/RPM-1 binding proteins. Moreover, we now provide an extensive array of citations to note that genetic enhancer interactions between different RPM-1/MYCBP2 binding proteins is well established. The reviewer has nicely highlighted for us that we handled the C. elegans genetics in too cursory a fashion in our original manuscript. We appreciate this being noted and have now aimed to make this substantially clearer. We hope the reviewer agrees that our revised C. elegans section accomplishes this goal.

Furthermore, we extensively revised the text of the Results to emphasize a key point: our observation that axon termination defects are not suppressed in vab-1; fsn-1 and vab-1; rpm-1 double mutants excludes the possibility that the VAB-1 Eph receptor is a substrate that is inhibited or degraded by the RPM-1/FSN-1 ubiquitin ligase complex. If the VAB-1 Eph receptor were ubiquitinated and degraded by the RPM-1/FSN-1 complex, we would have observed a suppression of phenotype in vab-1; rpm-1 double mutants. The precedent for this genetic relationship between the RPM-1 ubiquitin ligase and its substrates that are degraded has been established by several prior studies (PMID: 15707898; PMID: 31676756; PMID: 35421092). We now more clearly note that the absence of genetic suppression in vab-1; rpm-1 double mutants and vab-1; fsn-1 double mutants is consistent with the non-canonical stabilizing role of MYCBP2 on EPHB2 that was observed in our biochemical experiments with mammalian cells.

We also adjusted the text of the manuscript to stress that we are testing genetic interactions between the VAB-1 Eph receptor and known RPM-1 binding proteins. This is a key point, as genetic enhancer interactions are consistent with the Eph receptor functioning in the RPM-1 signaling network. This concept has been well established for RPM-1 binding proteins as now noted in our revised text with an extensive number of additional citations to published work.

Based on the above arguments, we respectfully disagree with the reviewer that our C. elegans data should be removed from the paper. To re-iterate, we are not trying to evaluate specificity for MYCBP2 and EPHB2 in C. elegans. Rather, our goals are twofold: 1) To ask whether there is an evolutionarily conserved functional genetic link between Eph receptors and known RPM-1 binding proteins. 2) To provide further in vivo genetic evidence invalidating the hypothesis that Ephrin receptors could be ubiquitination substrates that are inhibited/degraded by MYCBP2.

Text edits reflecting these points are in the abstract, the C. elegans results section starting on line 411, and the discussion on lines 499, 502-504 and 541.

The cellular responses are not robust and the effects of MYCBP2 KO - although significant - are minor in most cases. But I don't think more experiments will help here.

We interpret the comment about the robustness to mean that the extent to which a given cellular response is affected by the loss of MYCBP2 is minor. First, the cellular responses themselves are typical of previous studies and depend on the cellular biology underlying them. For example, a growth collapse of ~50-60% over a background of 10% (Fig. 7) is typical for these sorts of assays (PMID: 37369692; PMID: 33972524; PMID: 17785182). A decrease of cell area by ~25% (Fig. 3) is quite substantial if one considers how much of a cell’s volume is taken up by the nucleus and organelles. Second, the phenotypes elicited by the loss of MYCBP2 are likely brought on by a decrease in EphB2 protein levels, but not its complete absence, as suggested by our biochemical experiment. Given that EphB2 complete loss only affects the cellular responses to a limited extent, the minor effects are not a surprise (e.g. for GC collapse: PMID: 23143520). Nevertheless, the subtle changes in cellular phenotypes, elicited by EPHB2 signaling are often sufficient to achieve proper cell positioning and cell response to guidance cues. For instance, regulation of the growth cone collapse of the outgrowing axons requires delicate changes that are dynamic and temporal.

Minor:

Fig 1C - EPHA3 and EPHB2 seem to run in different sizes, is this the case? In 2A they run at the same size.

We believe this size discrepancy is due to different percentages of SDS-PAGE gels used to resolve proteins. In Fig. 1C, we used a 6% gel for a Western blot analysis of both EPHA3/-B2-FLAG (~130 kDa) and MYCBP2 (~510 kDa). In Fig. 2A however, we performed Western blot analysis using 10% resolving gel to separate and detect EPHA3/-B2-FLAG along with MYC-FBXO45 (~30 kDa). We have reviewed the results obtained from additional biological replicates of this experiment, and observed a similar pattern in gel migration of EPHA3/-B2-FLAG across all replicates.

Fig1F - I can't trust the MYCBP2 blot.

Indeed, the MYCBP2-EPHB2 co-IP with endogenous proteins was not convincing. We now repeated this experiment using rat cortical neurons, and the results replace the previous Fig. 1F panel as mentioned on line 158.

In Fig2b the authors claim that there is enhancement in the binding of MYCBP2 and EPHB2 upon FBXO45 expression. For this type of statement quantification is required.

The quantification is now included in Fig. 2C and its significance is mentioned on line 180. Our conclusion about the enhancement stands.

Fig2G - it remained unclear to me where the binding site to MYCBP2 is, how long is the cytoplasmic tail in the DeltaICD protein?

Based on our experimental observations from Fig. 2E-H, we concluded that the fragment encompassing the extracellular domain(s) and/or transmembrane (TM) domain of EPHB2 is necessary for the protein complex formation with MYCBP2. We would like to accentuate that the EPHB2-MYCBP2 interaction might not be direct, and might involve other transmembrane protein(s) acting as a scaffold for EPHB2 and MYCBP2 binding. We did not pursue experiments to determine the exact region of the extracellular-TM portion of EPHB2 that is required for the interaction with MYCBP2.

The cytoplasmic tail in ΔICD protein consists of 25 aa of the N-terminal fragment of EPHB2 juxtamembrane (JM) region, which is adjacent to the TM helix, and followed by the 8 aa FLAG tag (EPHB2 ΔICD domain composition: extracellular domains – TM domain – 25 aa fragment of JM region – FLAG). We have determined the TM and JM sequences based on Hedger et al. (PMID: 25779975) and included the N-terminal portion of the JM region to facilitate proper ΔICD protein localization within the plasma membrane (PMID: 35793621). We modified the schematic in Fig. 2G to better visualise the EPHB2 truncations and now provide information on their size in the figure legend.

Always good to have a model of how all these proteins work together.

While we acknowledge that this would be helpful, we do not have a clear answer on how the EPHB2-MYCBP2 complex formation occurs. This requires further elucidation of the putative proteins involved in this ternary complex or testing the possibility that a MYCBP2 fragment is extruded extracellularly. Without these experiments there are too many possibilities to summarise into a clear model figure. We thus did not make any edits regarding these possibilities in the section starting on line 195.

Reviewer #2 (Recommendations For The Authors):

Overall, the experiments are classical experiments of co-immunoprecipitations, swapping experiments, collapse assays, and stripe assays which all are well carried out and are convincing.

Thank you for your encouraging comments.

Controls for the stripe assay may include Fc / Fc stripe assays.

We have performed these control experiments and now include their quantifications in the results sectioning concerning Fig. 3, starting on line 249, and those concerning Fig. 6 on line 381.

It is not clear to me why SD and not SEM has been used here for presentations.

Standard deviation (SD) measures the dispersion of a dataset relative to its mean. The standard error of the mean (SEM) measures how much discrepancy is likely in a sample’s mean compared with the population mean. Thus, SEM includes a statistical inference about the sampling distribution while SD is a less “processed” measurement that by definition is larger than SEM. SEM might make the data look less dispersed and many journals encourage the use of SD in bar graphs (PMID: 16223828).

Fig 7A: it is rather difficult to see 'branches' in Fig. 7A, better pictures and close-ups should be provided. How are branches defined? This piece of work needs more attention.

To remedy this shortcoming, we now provide inverted images with GFP signal in dark pixels overlaid on Fc (white) / eB2 (pink) stripes next to the original images.

Reviewer #3 (Recommendations For The Authors):

1. My most important suggestion to the authors would be to more carefully describe the results and their interpretation of the results. Sometimes, the distinction is not clear.

We modified the text throughout the manuscript to address this.

1. There are several cases, when the authors report on trends that are not statistically significant (1D, for example), or report no change, when it is clear that the addition of one more sample could have dramatically made a difference (4M - see point 12).

We agree that some of the nonsignificant differences could become significant if we added more Ns. But we prefer not to move our experimental design towards N-chasing and p-hacking (PMID: 25768323). The number of biological replicates is normally pre-determined before the onset of the experiment. Of course, some replicates can be discarded if there is a valid reason, such as a technical issue with the experiment or a positive control not working but this is not relevant for the dataset we have provided.

1. Data in 1F is very difficult to interpret.

As in response to Reviewer #1: Indeed, the MYCBP2-EPHB2 co-IP with endogenous proteins was not convincing. We now repeated this experiment using rat cortical neurons, and the improved results are in revised Fig. 1F.

1. Figure 2 puts Figure 1 in a strange perspective. If I understand correctly, fig 2 claims that EPHB2 interaction with MYCBP2 depends on FBXO45 - if that is the case then how does the binding in Figure 1 occur?

Indeed, we propose that the EPHB2-MYCBP2 interaction depends on FBXO45. In Fig. 2, we reveal that FBXO45 enhances the formation of the EPHB2-MYCBP2 complex. Thus, we suspect that the endogenous FBXO45 present in HeLa cells and neurons would mediate the interaction between EPHB2 and MYCBP2 in Fig. 1 experiments. We were unable to show this by Western blotting due to lack of reliable commercial antibodies against FBXO45, the complex containing endogenous FBXO45 and EPHB2 is also implied by our AP-MS data (Fig. 1B) and published databases.

1. I am still trying to wrap my mind around the results in 2G-H. So do MYCBP2 and FBXO45 bind the extracellular domain of EPHBP2? What does that mean?

(see also our response to Reviewer #1, end of their section) Based on our experimental observations from Fig. 2G-H, we conclude that the fragment encompassing the extracellular domain(s) and/or transmembrane domain of EPHB2 is necessary for the protein complex formation with MYCBP2 and FBXO45. Although there is a possibility that MYCBP2 directly binds the extracellular portion of EPHB2, we have not formally tested this hypothesis. MYCBP2 has been previously shown to interact with the extracellular portion of transmembrane N-cadherin (CDH2) via BioID proximity labeling and AP-MS proteomics approaches (PMID: 32341084).

Considering the results in Fig. 2A-B, we suspect that EPHB2-MYCBP2 interaction is indirect, as FBXO45 enhances this association. Secretion of FBXO45 and direct binding of FBXO45 to the extracellular cadherin (EC1-2) domains of N-cadherin has been documented (PMID: 25143387; PMID: 32341084). Although, not tested, this is also a possibility for EPHB2-FBXO45 mode of interaction. Nevertheless, we also cannot rule out the possibility that an unknown transmembrane protein binds EPHB2 extracellularly and the same unknown protein binds MYCBP2/FBXO45 intracellularly. Resolving this model is beyond the scope of this study and will require us to pursue extensive new lines of investigation.

1. I don't understand the stable Hela cell line CRISPR - is this a stable MYCBP2 deletion? In which case why is there only a reduction, not complete elimination of the protein? Or, is this a stable integration of a plasmid generating gRNA against MYCBP2? In which case, I would expect a homozygous null to emerge at some point. In any case, this is not well explained.

These lines are not derived from single cells infected with the CRISPR sgRNA-carrying viruses, therefore they are not clonal and probably contain some cells that express normal levels of MYCBP2, hence its detection on a Western. This is now clarified starting on line 221 and on line 608.

1. In 3C - is this the right statistical analysis?? I would say you want to claim the different effect of the control +/- eB2 compared to the effect in the mutant +/- eB2. Still should be significant but I think a more correct analysis.

We now include this comparison in Fig. 3C as well in the results section starting on line 234.

1. The robustness of the assay in Figure 3D is underwhelming – how was the area measured?

This is a live imaging experiment. Fig. 3D plots cell area at 60 minutes after ephrin-B2 addition as a fraction of the same cell’s area at 0 minutes (ephrin-B2 addition). For control cells that is a decrease of ~25%. If one considers that a cell’s nucleus and organelles like the Golgi Apparatus take up most of its volume, the magnitude is not that surprising.

1. Figure 3F – did you try to plot the relative area of overlap divided by the total cellular area? You might get a more striking phenotype. Also – claiming that this confirms that MYCBP2 is REQUIRED for EPHB2 function is a bit overstated, especially given that we don’t know (do you?) the EPHB2 mutant phenotype in this assay.

We preferred to stay with the original method of image quantification which we use for other assays. With respect to the requirement of MYCBP2 for EPHB2 function in the stripe assay, our logic is rooted in the observation that native HeLa cells do not respond to ephrin-B2 stripes (45.46 ± 7.62% of cells on eB2 stripes v. Fc; data not shown). When they are transfected with EPHB2 expression plasmids they do, therefore we assume that EPHB2 expression endows them with a sensitivity to eB2 stripes. A loss of MYCBP2 attenuates this sensitivity. We clarified this starting on line 246 and on line 251.

1. I didn't quite get the difference between 4A and 4B.

We apologize for the confusion. In Fig 4A, we used a stable HeLa cell line that has tetracycline-inducible expression of EPHB2-FLAG. Using these cells, we subsequently generated CTRLCRISPR or MYCBP2CRISPR cells. In these cells we then induced EPHB2 expression with tetracycline and observed that deletion of MYCBP2 resulted in the reduction of EPHB2 protein levels. To confirm this observation and to rule out the possibility that EPHB2 protein reduction is an effect of the CRISPR lines generation, we tested whereas MYCBP2 deletion reduces EPHB2, which has been transiently overexpressed (Fig. 4B). We hence conclude that loss of MYCBP2 decreases EPHB2 that was either expressed from a stable locus (Fig. 4A) or from transient transfection (Fig. 4B). We modified the Results section starting on line 262 to make this point clear.

1. The entire link to lysosomal degradation should be strengthened. Perhaps I am confused, but if the reduced EPHB2 levels in MYCBP2 mutant cells result from impaired lysosomal degradation then inhibiting the lys-deg should bring the protein levels back to normal (i.e. CRISPR control) - no? As currently presented, I do not understand nor do I think the claim is strongly supported by the data.

Before treatment with inhibitors, EPHB2 levels in MYCBP2CRISPR cells are already 40% lower than they are in CTRLCRISPR cells and in all our attempts, inhibitors can only rescue/restore EPHB2 in MYCBP2CRISPR cells to a level that is lower than in CTRLCRISPR cells. But this restoration is greater in MYCBP2CRISPR than in MYCBP2CTRL cells (BafA1: 19% increase in CTRL cells and 40% in MYCBP2CRISPR cells; CoQ: 10% comparing to 35%). This indicates that EPHB2 degradation through the lysosomal pathway in MYCBP2CRISPR cells is stronger, explaining why EPHB2 degradation is promoted in MYCBP2CRISPR cells, compatible with reduced EPHB2 levels and enhanced EPHB2 ubiquitination.

1. 4M, O - reporting ns based on these data seems a bit strange to me... Add one point and it will be strongly significant.

See our response to point (2), above. We prefer not to invoke potential p-hacking.

1. 7d - so what are you claiming? That the cellular response to eB1 but not eB2 is affected by the addition of FBD1? this is almost the opposite of what you wrote in the text...

We treated the cells with two different ephrin-B ligands to make a stronger conclusion. When using ephrin-B1, growth cone collapse in FBD1 WT is not significant comparing to Fc treatment. When using ephrin-B2, growth cone collapse in FBD1 WT is not as significant as it is in FBD1 mut group (*** versus **). We interpret this as meaning that the EPHB2-mediated growth cone collapse to both ligands is dampened, when we disrupt the EPHB2-MYCBP2 association. The difference between these two ligands might be due to their different affinities for the receptor or signalling kinetics.

1. By far the weakest link in this paper is the worm part. I think it's a pity because strengthening this would affect the significance of the finding. First, the authors mention new genes without introducing their relationship to the signaling pathway tested. Second, the textual logics should be strengthened. Finally and most importantly, when the difference between the phenotypic severity is so strong (vab-1 and rpm-1) then I think it's impossible to say anything from the double mutant.

We appreciate the reviewer noting that they appreciate the value and importance of the C. elegans model. The goals of our C. elegans experiments were twofold:

1. To evaluate genetic interactions between the VAB-1 Eph receptor and known RPM-1 binding proteins. This was not clearly explained in the original manuscript nor was the published precedent for these types of genetic enhancer experiments provided. We have now rectified this by substantially revising the text of the Results C. elegans section starting on line 431 and by adding several citations.

2. Our C. elegans genetics confirmed that the VAB-1 Eph receptor is not inhibited/degraded by the RPM-1/MYCBP2 ubiquitin ligase complex. We have now revised the text to draw this point out more clearly.

To further address the reviewer’s concerns, we have added a new schematic (Fig. 8A) to show the relationship between the RPM-1 and the RPM-1 binding proteins (FSN-1/FBXO45 and GLO-4/SERGEF) we are testing. We chose FSN-1 because it is part of the RPM-1 ubiquitin ligase complex and we chose GLO-4 because it functions outside the context of RPM-1 ubiquitin ligase signaling via the GLO-1 Rab GTPase to influence late endosomal/lysosomal biogenesis.

Regarding the reviewer’s concern that different penetrance/frequency of defects between rpm-1 mutants and vab-1 mutants means outcomes with vab-1; rpm-1 double mutants cannot be interpreted. We respectfully disagree. An extensive number of published studies have demonstrated that RPM-1 binding proteins have milder phenotypes than rpm-1 mutants and display genetic enhancer effects as double mutants with one another (PMID:17698012, PMID: 22357847, PMID: 25010424, PMID: 24810406). We now make this point much more clearly. While the frequency of axon termination defects in rpm-1 mutants is high it is not completely saturated as the defect is not 100%. Moreover, a major point of the vab-1; rpm-1 double mutants is that they do not have a significant reduction in phenotypic penetrance/frequency. Thus, our system is fully capable of resolving genetic suppression, which did not occur. We now make this point much more carefully and clearly.

To further address the reviewer’s concern, we have softened language about the VAB-1/Eph receptor functioning in the same pathway as RPM-1 throughout the manuscript. While we think this is still the case, because the frequency of axon termination defects is not fully saturated in rpm-1 mutants and defects could potentially become more severe (i.e. the hook might occur closer to the head of the animal rather than in the midbody). Nonetheless, this is not a critical point and we think it is more important to be clear about the two major goals and objectives of our C. elegans experiments. We hope the reviewer agrees that our rationale, logic and conclusions are more clearly and accurately drawn in the revised paper.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Eph receptor tyrosine kinases participate in a variety of normal and pathogenic processes during development and throughout adulthood. This versatility is likely facilitated by the ability of Eph receptors to signal through diverse cellular signalling pathways: primarily by controlling cytoskeletal dynamics, but also by regulating cellular growth, proliferation, and survival. Despite many proteins linked to these signalling pathways interacting with Eph receptors, the specific mechanisms behind such links and their coordination remain to be elucidated. In a proteomics screen for novel EPHB2 multi-effector proteins, we identified human MYC binding protein 2 (MYCBP2 or PAM or Phr1). MYCBP2 is a large signalling hub involved in diverse processes such as neuronal connectivity, synaptic growth, cell division, neuronal survival, and protein ubiquitination. Our biochemical experiments demonstrate that the formation of a complex containing EPHB2 and MYCBP2 is facilitated by FBXO45, a protein known to select substrates for MYCBP2 ubiquitin ligase activity. Formation of the MYCBP2-EPHB2 complex does not require EPHB2 tyrosine kinase activity and is destabilised by binding of ephrin-B ligands, suggesting that the MYCBP2-EPHB2 association is a prelude to EPHB2 signalling. Paradoxically, the loss of MYCBP2 results in increased ubiquitination of EPHB2 and a decrease of its protein levels suggesting that MYCBP2 stabilises EPHB2. Commensurate with this effect, our cellular experiments reveal that MYCBP2 is essential for efficient EPHB2 signalling responses in cell lines and primary neurons. Finally, our genetic studies in C. elegans provide in vivo evidence that the ephrin receptor VAB-1 functions within the MYCBP2 signalling network. Together, our results align with the similarity of neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function, and our findings couple EPHB2 to a signaling effector controlling diverse cellular functions.

Revised version: Eph receptor tyrosine kinases participate in a variety of normal and pathogenic processes during development and throughout adulthood. This versatility is likely facilitated by the ability of Eph receptors to signal through diverse cellular signalling pathways: primarily by controlling cytoskeletal dynamics, but also by regulating cellular growth, proliferation, and survival. Despite many proteins linked to these signalling pathways interacting with Eph receptors, the specific mechanisms behind such links and their coordination remain to be elucidated. In a proteomics screen for novel EPHB2 multi-effector proteins, we identified human MYC binding protein 2 (MYCBP2 or PAM or Phr1). MYCBP2 is a large signalling hub involved in diverse processes such as neuronal connectivity, synaptic growth, cell division, neuronal survival, and protein ubiquitination. Our biochemical experiments demonstrate that the formation of a complex containing EPHB2 and MYCBP2 is facilitated by FBXO45, a protein known to select substrates for MYCBP2 ubiquitin ligase activity. Formation of the MYCBP2-EPHB2 complex does not require EPHB2 tyrosine kinase activity and is destabilised by binding of ephrin-B ligands, suggesting that the MYCBP2-EPHB2 association is a prelude to EPHB2 signalling. Paradoxically, the loss of MYCBP2 results in increased ubiquitination of EPHB2 and a decrease of its protein levels suggesting that MYCBP2 stabilises EPHB2. Commensurate with this effect, our cellular experiments reveal that MYCBP2 is essential for efficient EPHB2 signalling responses in cell lines and primary neurons. Finally, our genetic studies in C. elegans provide in vivo evidence that the ephrin receptor VAB-1 displays genetic interactions with known MYCBP2 binding proteins. Together, our results align with the similarity of neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function, and couple EPHB2 to a signaling effector that controls diverse cellular functions.

Introduction

Earlier version: Eph receptor tyrosine kinases and their membrane-tethered ligands, the ephrins, elicit short distance cell-cell signals that regulate many biological processes. Ephrin-Eph signalling primarily impacts the cytoskeleton with the immobilization of highly dynamic axonal growth cones being a classic example. Other processes that involve changes in transcription, growth, and survival such as angiogenesis, synaptic plasticity, stem cell fate, tumorigenesis and neurodegeneration also involve the Eph/ephrin system. Many proteins are postulated to couple Eph receptors to different intracellular effectors, but the molecular logic of this diversity remains fragmented (Bush, 2022; Kania and Klein, 2016).

Revised version: Eph receptor tyrosine kinases and their membrane-tethered ligands, the ephrins, elicit short distance cell-cell signals that regulate many biological processes. Ephrin-Eph signalling primarily impacts the cytoskeleton with the immobilization of highly dynamic axonal growth cones being a classic example. Other processes that involve changes in transcription, growth, and survival such as angiogenesis, synaptic plasticity, stem cell fate, tumorigenesis and neurodegeneration also involve the Eph/ephrin system. Many proteins are postulated to couple Eph receptors to different intracellular effectors, but the molecular logic of this diversity remains fragmented (Kania and Klein, 2016;Bush, 2022).

Earlier version: EphB subfamily members preferentially bind transmembrane ephrin-Bs and although both molecules participate in bidirectional signalling, ephrin-B activation of EphB signalling cascades is more thoroughly studied (Gale et al., 1996; Mellitzer et al., 1999). To elicit robust Eph receptor forward signalling, ephrins multimerise in signalling clusters by intercalating with Ephs on a signal-recipient cell with array size correlating with signal amplitude (Kullander et al., 2001; Schaupp et al., 2014). Signalling initiation involves the activation of the receptor tyrosine kinase and phosphorylation of tyrosines proximal to the EphB transmembrane domain (Binns et al., 2000; Soskis et al., 2012). Eph-evoked cytoskeletal effects such as cell contraction and growth cone collapse result from changes in small GTPase activity modulated by EphB Guanine nucleotide exchange factors (GEFs) and GTPase activating proteins (GAPs) (Margolis et al., 2010; Shi et al., 2007). Eph signalling has also been linked to fundamental cellular pathways such as the Ras-MAPK pathway, mTOR-regulated protein synthesis, cell division and survival (Bush and Soriano, 2010; Depaepe et al., 2005; Fawal et al., 2018; Genander et al., 2009; Nie et al., 2010). Eph forward signalling eventually leads to their internalisation and either recycling or degradation via endosome/lysosome and ubiquitination/proteasome pathways (Okumura et al., 2017; Zimmer et al., 2003). While identification of a growing number of proteins interacting with Eph receptors has moved the field forward, we have yet to clarify the question of how Eph receptors activate various fundamental cellular processes, often within the same cell type.

Revised version: EphB subfamily members preferentially bind transmembrane ephrin-Bs and although both molecules participate in bidirectional signalling, ephrin-B activation of EphB signalling cascades is more thoroughly studied (Gale et al., 1996;Mellitzer et al., 1999). To elicit robust Eph receptor forward signalling, ephrins multimerise in signalling clusters by intercalating with Ephs on a signal-recipient cell with array size correlating with signal amplitude (Kullander et al., 2001;Schaupp et al., 2014). Signalling initiation involves the activation of the receptor tyrosine kinase and phosphorylation of tyrosines proximal to the EphB transmembrane domain (Soskis et al., 2012;Binns et al., 2000). Eph-evoked cytoskeletal effects such as cell contraction and growth cone collapse result from changes in small GTPase activity modulated by EphB Guanine nucleotide exchange factors (GEFs) and GTPase activating proteins (GAPs) (Margolis et al., 2010;Shi et al., 2007). Eph signalling has also been linked to fundamental cellular pathways such as the Ras-MAPK pathway, mTOR-regulated protein synthesis, cell division and survival (Bush and Soriano, 2010;Nie et al., 2010;Fawal et al., 2018;Genander et al., 2009;Depaepe et al., 2005). Eph forward signalling eventually leads to their internalisation and either recycling or degradation via endosome/lysosome and ubiquitination/proteasome pathways (Zimmer et al., 2003;Okumura et al., 2017). While identification of a growing number of proteins interacting with Eph receptors has moved the field forward, we have yet to clarify the question of how Eph receptors activate various fundamental cellular processes, often within the same cell type.

Earlier version: Myc-binding protein 2 (MYCBP2), also known as Highwire, RPM-1, Protein Associated with Myc (PAM) and Phr1, is a large signalling hub that regulates cytoskeletal dynamics, neuronal development, and axonal degeneration (Guo et al., 1998, Grill et al., 2016; Virdee, 2022). It has an atypical RING ubiquitin ligase activity that inhibits the p38/MAP and JNK kinase pathways thereby regulating cytoskeletal dynamics underlying axonal development and synaptic growth (Borgen et al., 2017; Collins et al., 2006; Lewcock et al., 2007; Nakata et al., 2005; Pao et al., 2018; Wan et al., 2000). MYCBP2 further regulates the Tuberin Sclerosis Complex linked to cell growth (TSC; (Han et al., 2012)), initiation of autophagy via ULK (Crawley et al., 2019) and NMNAT2-regulated neuronal survival and axonal degeneration (Babetto et al., 2013; Xiong et al., 2012). Biochemical mapping has shown that human MYCBP2 and its C. elegans ortholog RPM-1 rely upon the FBD1 domain to bind the F-box protein FBXO45 that acts as a ubiquitination substrate selector (Desbois et al., 2018); (Sharma et al., 2014). Intriguingly, the neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function are similar, raising the possibility that these two molecules could function in the same pathway (Dalva et al., 2000; Henkemeyer et al., 1996; Lewcock et al., 2007). Nonetheless, a biochemical or genetic interaction between MYCBP2 and EPHB2 has not been demonstrated in any system.

Revised version: Myc-binding protein 2 (MYCBP2), also known as Protein Associated with Myc (PAM) and Highwire, RPM-1, or Phr1 in different species, is a large signalling hub that regulates cytoskeletal dynamics, neuronal development, and axonal degeneration (Guo et al., 1998;Grill et al., 2016;Virdee, 2022). It has an atypical RING ubiquitin ligase activity that inhibits the p38/MAP and JNK kinase pathways thereby regulating cytoskeletal dynamics underlying axonal development and synaptic growth (Nakata et al., 2005;Collins et al., 2006;Pao et al., 2018;Wan et al., 2000;Lewcock et al., 2007;Borgen et al., 2017). MYCBP2 further regulates the Tuberin Sclerosis Complex linked to cell growth (Han et al., 2012), initiation of autophagy via ULK (Crawley et al., 2019) and NMNAT2-regulated neuronal survival and axonal degeneration (Babetto et al., 2013;Xiong et al., 2012). Biochemical mapping has shown that human MYCBP2 and its C. elegans ortholog RPM-1 rely upon the FBD1 domain to bind the F-box protein FBXO45 that acts as a ubiquitination substrate selector (Desbois et al., 2018;Sharma et al., 2014). Intriguingly, the neurodevelopmental phenotypes caused by MYCBP2 and EPHB2 loss of function are similar, raising the possibility that these two molecules could function in the same pathway (Henkemeyer et al., 1996;Lewcock et al., 2007;Dalva et al., 2000). Nonetheless, a biochemical or genetic interaction between MYCBP2 and EPHB2 has not been demonstrated in any system.

Earlier version: To shed light on how EphB receptors fulfil their multitude of functions, we used mass spectrometry (MS)-based proteomics to identify multi-effector proteins that bind EPHB2. One of our proteomic hits was MYCBP2, which we demonstrated forms a complex with EPHB2 using a combination of biochemical and cellular assays. Furthermore, we show that this interaction is required for efficient EPHB2 signalling in cell lines and primary neurons. Consistent with these findings, we observed in vivo genetic interactions between the Eph receptor, VAB-1, and the RPM-1 signalling network in the nervous system of C. elegans. Our collective results indicate that the relationship between EPHB2 and MYCBP2 does not appear to involve the ubiquitin ligase activity of MYCBP2, and raise the possibility that MYCBP2 links EPHB2 to diverse fundamental cellular functions.

Revised version: To shed light on how EphB receptors fulfil their multitude of functions, we used mass spectrometry (MS)-based proteomics to identify multi-effector proteins that bind EPHB2. One of our proteomic hits was MYCBP2, which we demonstrated forms a complex with EPHB2 using a combination of biochemical and cellular assays. Furthermore, we show that this interaction is required for efficient EPHB2 signalling in cell lines and primary neurons. Consistent with these findings, we observed in vivo genetic interactions in C. elegans between the Eph receptor, VAB-1, and known RPM-1 binding proteins. Our collective results indicate that the relationship between EPHB2 and MYCBP2 does not appear to involve the ubiquitin ligase activity of MYCBP2, and raise the possibility that MYCBP2 links EPHB2 to diverse fundamental cellular functions.

Results

Proteomics identifies MYCBP2 as a putative EPHB2-interacting protein

To understand the molecular logic underlying EPHB2 signalling diversity, we performed affinity purification coupled to MS (AP-MS) in order to identify EPHB2-interacting proteins, and prioritised those known to be signalling hubs. We used a stable HeLa cell line with tetracycline-inducible expression of BirA-linked EPHB2-FLAG, that we previously used to study EPHB2 signalling (Lahaie et al., 2019). To identify ephrin ligand-dependent EPHB2 protein complexes, we stimulated EPHB2-FLAG-overexpressing cells with pre-clustered Fc control or ephrinB2-Fc (eB2-Fc). We then harvested and lysed the cells, performed anti-FLAG immunoprecipitation, and used mass spectrometry to identify EPHB2 protein complexes (Fig. 1A).

MS-proteomics and biochemistry in HeLa cells identifies MYCBP2 as EPHB2 binding protein.

Earlier version: A) Schematic of EPHB2 affinity purification coupled to mass spectrometry (AP-MS) workflow. B) Scatter plot of AP-MS data showing known and putative EPHB2 binding proteins, including MYCBP2. Y axis and X axes represent the average spectral counts of the identified protein hits in the EPHB2 protein complexes from cells stimulated with Fc control or ephrin-B2 (eB2-Fc), respectively. C) In HEK 293T cells, endogenous MYCBP2 is pulled down by transiently overexpressed EPHB2-FLAG but not by EPHA3-FLAG. D) In EPHB2-FLAG stable HeLa cell line, ephrin-B stimulation reduces the interaction between MYCBP2 and EPHB2. E) Quantification of MYCBP2-EPHB2 association intensity after Fc, ephrin-B1 (eB2-Fc) or ephrin-B2 (eB2-Fc) treatment (eB1-Fc, p=0.1365; eB2-Fc, p=0.0002; one-sample t-test). Error bars represent standard deviation (SD). F) Representative image of MYCBP2 pull down with anti-EPHB2 or IgG control antibodies from whole rat cerebral tissue lysates. Asterisk indicates MYCBP2. G) Representative images from western blot analysis of endogenous MYCBP2 following IP of GFP-EPHB2 wild-type (WT) or its kinase dead (KD) counterpart.

Revised version: A) Schematic of EPHB2 affinity purification coupled to mass spectrometry (AP-MS) workflow. B) Scatter plot of AP-MS data showing known and putative EPHB2 binding proteins, including MYCBP2. Y axis and X axes represent the average spectral counts of the identified protein hits in the EPHB2 protein complexes from cells stimulated with Fc control or ephrin-B2 (eB2-Fc), respectively. C) In HEK 293T cells, endogenous MYCBP2 is pulled down by transiently overexpressed EPHB2-FLAG but not by EPHA3-FLAG. D) In EPHB2-FLAG stable HeLa cell line, ephrin-B stimulation reduces the interaction between MYCBP2 and EPHB2. E) Quantification of MYCBP2-EPHB2 association intensity after Fc, ephrin-B1 (eB2-Fc) or ephrin-B2 (eB2-Fc) treatment (eB1-Fc, p=0.1365; eB2-Fc, p=0.0002; one-sample t-test). Error bars represent standard deviation (SD). F) Representative image of MYCBP2 pull down with anti-EPHB2 or IgG control antibodies from rat cortical neurons. Asterisk indicates MYCBP2. G) Representative images from western blot analysis of endogenous MYCBP2 following IP of GFP-EPHB2 wild-type (WT) or its kinase dead (KD) counterpart.

To identify EPHB2-specific interactions and remove background contaminants, we performed Significance Analysis of INTeractome (SAINTexpress analysis (Teo et al., 2014)) using our EPHB2-related controls (Lahaie et al., 2019). To better visualize the changes in putative EPHB2 binding partners, we compared the average spectral counts of the identified proteins using ProHits-viz tool (Knight et al., 2017). Comparison of Fc and ephrin-B2-treated samples did not yield any significant differences. Since we applied the ligand and collected the samples on a time scale comparable to known EphB signalling dynamics, this limitation could be potentially due to autoactivation of Eph receptors when they are overexpressed, obscuring some ligand-dependent effects (Lackmann et al., 1998). However, the resulting scatter plot confirmed the presence of several known EPHB2 interactors, such as FYN and YES1 Src kinases and members of the Eph receptor family (Fig. 1B) (Banerjee et al., 2022). One of the most prominent, novel hits was the E3 ubiquitin ligase and signalling hub protein MYCBP2, and its binding partner FBXO45. FBXO45 was previously identified as a putative EPHB2 interacting protein in large-scale, cell-based interactome studies (Huttlin et al., 2021;Salokas et al., 2022).

Biochemical validation of MYCBP2 binding to EPHB2

Earlier version: To confirm that MYCBP2 can indeed form a molecular complex with EPHB2, we tested whether endogenous MYCBP2 co-immunoprecipitates (co-IP) with FLAG-tagged EPHB2 in HEK 293T cells. Based on differences in EPHB2 and EPHA3 interactomes, we reasoned that EPHA3 may serve as a negative control for the EPHB2-MYCBP2 association (Huttlin et al., 2021). We found that MYCBP2 coprecipitated with affinity-purified EPHB2-FLAG, but not EPHA3 (Fig. 1C). This confirmed MYCBP2 binding to EPHB2 and suggested MYCPB2 displays EPH receptor subtype specificity. Importantly, the EPHB2-MYCBP2 interaction was reduced by 24.5% and 53% following ephrin-B1 and ephrin-B2 treatment respectively in HeLa EPHB2 cells, suggesting the involvement of MYCBP2 in ephrin-B:EPHB2 signalling (Fig. 1D, E, eB1-Fc, p=0.1365; eB2-Fc, p=0.0002). To test this interaction in vivo, we performed co-IP from rat brain, which further confirmed MCYBP2 association with EPHB2 (Fig. 1F).

Revised version: To confirm that MYCBP2 can indeed form a molecular complex with EPHB2, we tested whether endogenous MYCBP2 co-immunoprecipitates (co-IP) with FLAG-tagged EPHB2 in HEK 293T cells. Based on differences in EPHB2 and EPHA3 interactomes, we reasoned that EPHA3 may serve as a negative control for the EPHB2-MYCBP2 association (Huttlin et al., 2021). We found that MYCBP2 coprecipitated with affinity-purified EPHB2-FLAG, but not EPHA3 (Fig. 1C). This confirmed MYCBP2 binding to EPHB2 and suggested MYCPB2 displays EPH receptor subtype specificity. Importantly, the EPHB2-MYCBP2 interaction was reduced by 24.5% and 53% following ephrin-B1 and ephrin-B2 treatment respectively in HeLa EPHB2 cells, suggesting the involvement of MYCBP2 in ephrin-B:EPHB2 signalling (Fig. 1D, E, eB1-Fc, p=0.1365; eB2-Fc, p=0.0002). To test this interaction in vivo, we performed co-IP using dissociated rat cortical neurons, which further confirmed MCYBP2 association with EPHB2 (Fig. 1F).

We next asked whether the kinase activity of EPHB2 is required for the formation of the EPHB2-MYCBP2 complex. To test this, we expressed GFP-tagged wild type (WT) EPHB2 or EPHB2 with a kinase-dead mutation (KD) in HeLa cells. We found that MYCBP2 showed comparable coprecipitation with WT and KD EPHB2 (Fig. 1G). Thus, EPHB2 kinase activity is not required for the formation of MYCBP2 - EPHB2 complexes, suggesting that MYCBP2 association with EPHB2 may be a prelude to ephrin-B-evoked EPHB2 signalling.

FBXO45 enhances the association between MYCBP2 and EPHB2

Earlier version: Our proteomics screen for EPHB2 interactors also identified FBXO45, the F-box protein that forms a ubiquitin ligase complex with MYCBP2 (Saiga et al., 2009; Sharma et al., 2014). Thus, we initially reasoned that FBXO45 might perform a similar role in the formation of the MYCBP2-EPHB2 complex. We first tested whether FBXO45 binds to EPHB2 by co-expressing MYC-tagged FBXO45 with EPHB2-FLAG or EPHA3-FLAG in HEK293T cells. Co-IP revealed that FBXO45 can associate with EPHB2 but not EPHA3, suggesting that EPHB2, FBXO45 and MYCBP2 could form a ternary complex (Fig. 2A). To test this idea, we co-expressed EPHB2-FLAG and GFP-MYCBP2 in the presence or absence of MYC-FBXO45 and examined co-IP efficiency between EPHB2 and MYCBP2. Interestingly, FBXO45 enhanced the interaction between EPHB2 and MYCBP2 (Fig. 2B). Together, these data suggest that EPHB2 can form a complex with both MYCBP2 and FBXO45, and that FBXO45 increases the efficiency of MYCBP2-EPHB2 interaction.

Revised version: Our proteomics screen for EPHB2 interactors also identified FBXO45, the F-box protein that forms a ubiquitin ligase complex with MYCBP2 (Sharma et al., 2014;Saiga et al., 2009). Thus, we initially reasoned that FBXO45 might perform a similar role in the formation of the MYCBP2-EPHB2 complex. We first tested whether FBXO45 binds to EPHB2 by co-expressing MYC-tagged FBXO45 with EPHB2-FLAG or EPHA3-FLAG in HEK293T cells. Co-IP revealed that FBXO45 can associate with EPHB2 but not EPHA3, suggesting that EPHB2, FBXO45 and MYCBP2 could form a ternary complex (Fig. 2A). To test this idea, we co-expressed EPHB2-FLAG and GFP-MYCBP2 in the presence or absence of MYC-FBXO45 and examined co-IP efficiency between EPHB2 and MYCBP2. Interestingly, FBXO45 enhanced the interaction between EPHB2 and MYCBP2 (Fig. 2B, C, p=0.0068). Together, these data suggest that EPHB2 can form a complex with both MYCBP2 and FBXO45, and that FBXO45 increases the efficiency of MYCBP2-EPHB2 interaction.

Mapping binding regions for EPHB2-MYCBP2 reveals role of FBXO45 in this interaction.

Earlier version: A) Co-IP of EPHB2-FLAG with MYC-FBXO45 using transfected HEK293 cells. EPHB2 co-precipitates FBXO45, but EPHA3 does not. Asterisk indicates MYC-FBXO45. B) In HEK 293T cells, FBXO45 overexpression enhances EPHB2-MYCBP2 binding. Asterisk indicates MYC-FBXO45. C) Schematic representation of MYCBP2 N-terminal, Central and C-terminal fragments. D) Co-IP of EPHB2-FLAG with GFP-MYCBP2 fragments in HEK 293T cells. EPHB2 coprecipitates with MYCBP2 central fragment. Asterisks indicate GFP-MYCBP2 fragments. E) Schematic of chimeric domain swapping of EPHB2 (orange) and EPHA3 (grey). F) Co-IP of MYC-FBXO45 and endogenous MYCBP2 with EPHB2/EPHA3 domain swapped chimeras. G) Schematic representation of EPHB2 ΔECD (extracellular domain) and ΔICD (intracellular domain) truncations. H) Co-IP of endogenous MYCBP2 with EPHA3, EPHB2 and EPHB2 truncation mutants. ECD, extracellular domain; TM, transmembrane; JM, juxtamembrane; SAM, Sterile alpha motif; PBM, PDZ (PSD-95, Dlg1, Zo-1) binding motif.

Revised version: A) Co-IP of EPHB2-FLAG with MYC-FBXO45 using transfected HEK293 cells. EPHB2 co-precipitates FBXO45, but EPHA3 does not. Asterisk indicates MYC-FBXO45. B) In HEK 293T cells, FBXO45 overexpression enhances EPHB2-MYCBP2 binding. Asterisk indicates MYC-FBXO45. C) Quantification of the association intensity of MYCBP2 and EPHB2 upon FBXO45 overexpression (EPHB2, p=0.0068; EPHB2 vs EPHA3, p=0.0005; one sample t-test). Error bars represent SD. D) Schematic representation of MYCBP2 N-terminal, Central and C-terminal fragments. E) Co-IP of EPHB2-FLAG with GFP-MYCBP2 fragments in HEK 293T cells. EPHB2 coprecipitates with MYCBP2 central fragment. Asterisks indicate GFP-MYCBP2 fragments. F) Schematic of chimeric domain swapping of EPHB2 (orange) and EPHA3 (grey). G) Co-IP of MYC-FBXO45 and endogenous MYCBP2 with EPHB2/EPHA3 domain swapped chimeras. H) Schematic representation of EPHB2 ΔECD (extracellular domain, aa deletions of 19-530) and ΔICD (intracellular domain, aa deletions of 590-986) truncations. I) Co-IP of endogenous MYCBP2 with EPHA3, EPHB2 and EPHB2 truncation mutants. ECD, extracellular domain; TM, transmembrane; JM, juxtamembrane; SAM, Sterile alpha motif; PBM, PDZ (PSD-95, Dlg1, Zo-1) binding motif.

Biochemical mapping of EPHB2-MYCBP2 interaction

Earlier version: To identify the MYCBP2 protein domain(s) required for the formation of the ternary complex with EPHB2 and FBXO45, we co-expressed EPHB2-FLAG and three GFP-MYCBP2 fragments in HEK293T cells. Co-IP revealed that the central region of MYCBP2 was sufficient for binding with EPHB2 (Fig. 2C, D). In addition, co-expression of MYC-FBXO45 demonstrated that the association of the central domain of MYCBP2 with EPHB2 is enhanced by FBOX45 and is consistent with the presence of an FBXO45 binding site within this MYCBP2 fragment (Fig. 2D) (Sharma et al., 2014).

Revised version: To identify the MYCBP2 protein domain(s) required for the formation of the ternary complex with EPHB2 and FBXO45, we co-expressed EPHB2-FLAG and three GFP-MYCBP2 fragments in HEK293T cells. Co-IP revealed that the central region of MYCBP2 was sufficient for binding with EPHB2 (Fig. 2D, E). In addition, co-expression of MYC-FBXO45 demonstrated that the association of the central domain of MYCBP2 with EPHB2 is enhanced by FBOX45 and is consistent with the presence of an FBXO45 binding site within this MYCBP2 fragment (Fig. 2E) (Sharma et al., 2014).

Earlier version: To identify the domains of EPHB2 required for the formation of the tripartite complex, we took advantage of the observation that EPHA3 does not readily form a complex with MYCBP2 or FBXO45. Thus, we performed domain swapping experiments between EPHB2 and EPHA3 reasoning that placing EPHA3-specific sequences in EPHB2 would inhibit the formation of the tripartite complex. We constructed a series of FLAG-EPHB2/EPHA3 chimeras and determined whether they bound MYCBP2 in the presence of FBXO45 (Fig. 2E). Co-IP revealed that the ability of a particular chimera to associate with MYCBP2 was correlated with its association with FBXO45, in line with the MYCBP2-FBXO45 complex interacting with EPHB2. Surprisingly, we found that EPHB2-EPHA3 chimeras with an EPHA3 identity of intracellular juxtamembrane, kinase, SAM or PDZ binding domains retained the ability to associate with FBXO45 and MYCBP2, suggesting that the formation of the tripartite complex is driven by the extracellular domain and/or the transmembrane domain of EPHB2 (Fig. 2F). However, these results are also consistent with the possibility that the EPHB2 identity of the extracellular fragments could alter the conformation of the intracellular domains of EPHA3 identity, allowing the interaction with FBXO45-MYCBP2 to occur. To exclude this possibility, we created EPHB2 mutants lacking the intracellular or extracellular domains and tested their ability to complex with MYCBP2 and FBXO45 by co-IP (Fig. 2G). In these experiments only the deletion mutant lacking the intracellular domain retained its ability to form the tripartite complex. Collectively, these results argue that the combination of extracellular and transmembrane domains of EPHB2 are necessary and sufficient for formation of the MYCBP2–FBXO45–EPHB2 complex (Fig. 2H).

Revised version: To identify the domains of EPHB2 required for the formation of the tripartite complex, we took advantage of the observation that EPHA3 does not readily form a complex with MYCBP2 or FBXO45. Thus, we performed domain swapping experiments between EPHB2 and EPHA3 reasoning that placing EPHA3-specific sequences in EPHB2 would inhibit the formation of the tripartite complex. We constructed a series of FLAG-EPHB2/EPHA3 chimeras and determined whether they bound MYCBP2 in the presence of FBXO45 (Fig. 2F). Co-IP revealed that the ability of a particular chimera to associate with MYCBP2 was correlated with its association with FBXO45, in line with the MYCBP2-FBXO45 complex interacting with EPHB2. Surprisingly, we found that EPHB2-EPHA3 chimeras with an EPHA3 identity of intracellular juxtamembrane, kinase, SAM or PDZ binding domains retained the ability to associate with FBXO45 and MYCBP2, suggesting that the formation of the tripartite complex is driven by the extracellular domain and/or the transmembrane domain of EPHB2 (Fig. 2G). However, these results are also consistent with the possibility that the EPHB2 identity of the extracellular fragments could alter the conformation of the intracellular domains of EPHA3 identity, allowing the interaction with FBXO45-MYCBP2 to occur. To exclude this possibility, we created EPHB2 mutants lacking the intracellular or extracellular domains and tested their ability to complex with MYCBP2 and FBXO45 by co-IP (Fig. 2H). In these experiments only the deletion mutant lacking the intracellular domain retained its ability to form the tripartite complex. Collectively, these results argue that the combination of extracellular and transmembrane domains of EPHB2 are necessary and sufficient for formation of the MYCBP2–FBXO45–EPHB2 complex (Fig. 2I).

MYCBP2 is required for EPHB2-mediated cellular responses

Earlier version: Given the decrease in MYCBP2–EPHB2 association evoked by ephrin-B treatment (Fig. 1D, E), we next sought to determine whether MYCBP2 fulfils a specific function in ephrin-B:EPHB2 forward signalling. Thus, we used EPHB2-FLAG HeLa cells to generate a MYCBP2CRISPR HeLa cell line, which stably integrated CRISPR sgRNA targeting the MYCBP2 exon 6 using lentivirus (Desbois et al., 2018). Using EPHB2-FLAG HeLa cells carrying a stably integrated empty expression vector (CTRLCRISPR) as controls, we found that MYCBP2CRISPR led to a reduction in endogenous MYCBP2 protein levels (Fig. 3A).

Revised version: Given the decrease in MYCBP2–EPHB2 association evoked by ephrin-B treatment (Fig. 1D, E), we next sought to determine whether MYCBP2 fulfils a specific function in ephrin-B:EPHB2 forward signalling. Thus, we infected EPHB2-FLAG HeLa cells using lentivirus containing CRISPR sgRNA targeting the MYCBP2 exon 6, and pooled MYCBP2CRISPR cells after puromycin selection (Desbois et al., 2018). Using EPHB2-FLAG HeLa cells carrying a stably integrated empty expression vector (CTRLCRISPR) as controls, we found that MYCBP2CRISPR led to a reduction in endogenous MYCBP2 protein levels (Fig. 3A).

MYCBP2 CRISPR HeLa cells exhibit reduced ephrin-B2 evoked cell retraction and ephrin-B2 stripe avoidance.

Earlier version: A) MYCBP2 is reduced in HeLa MYCBP2CRISPR cells generated by sgRNA targeting MYCBP2 exon 6. B) Representative images of cell collapse assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells that were stimulated with ephrin-B2. Red arrows indicate rounded/collapsed cells. C) Quantification of collapsed cells. Statistical significance between CTRLCRISPR and MYCBP2CRISPR cells was determined using two-way ANOVA followed by Sidak’s multiple comparison test (Fc, p=0.9841; eB2-Fc, p=0.0017). D) Representative time-lapse sequences of CTRLCRISPR and MYCBP2CRISPR HeLa cells after ephrin-B2 treatment. E) Quantification of cell area reduction after 60 min exposure to ephrin-B2 (p=0.0268, two-tailed unpaired t test). F) Ephrin-B2 stripe assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells. Cells are visualized with Phalloidin 488 staining and nuclei are stained with DAPI (black, Fc stripes; red, ephrin-B2 stripes). G) Quantification of cells present on ephrin-B2 stripes (%). Statistical significance was determined using two-tailed unpaired t-test (p=0.0109). Error bars represent SD.

Revised version: A) MYCBP2 is reduced in HeLa MYCBP2CRISPR cells generated by sgRNA targeting MYCBP2 exon 6. B) Representative images of cell collapse assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells that were stimulated with ephrin-B2. Red arrows indicate rounded/collapsed cells. C) Quantification of collapsed cells. Statistical significance between CTRLCRISPR and MYCBP2CRISPR cells was determined using two-way ANOVA followed by Sidak’s multiple comparison test (CTRLCRISPR vs MYCBP2 CRISPR: Fc, p=0.9841; eB2-Fc, p=0.0017. Fc vs eB2-Fc: CTRLCRISPR, p<0.0001; MYCBP2 CRISPR, p=0.0115). D) Representative time-lapse sequences of CTRLCRISPR and MYCBP2CRISPR HeLa cells after ephrin-B2 treatment. E) Quantification of cell area reduction after 60 min exposure to ephrin-B2 (p=0.0268, two-tailed unpaired t test). F) Ephrin-B2 stripe assays using CTRLCRISPR or MYCBP2CRISPR HeLa cells. Cells are visualized with Phalloidin 488 staining and nuclei are stained with DAPI (black, Fc stripes; red, ephrin-B2 stripes). G) Quantification of cells present on ephrin-B2 stripes (%). Statistical significance was determined using two-tailed unpaired t-test (p=0.0109). Error bars represent SD.

Earlier version: To study the role of MYCBP2 in EphB signalling, we took advantage of an experimental paradigm in which exposure to ephrin-B2 evokes cytoskeletal contraction of HeLa cells expressing EPHB2 (Lahaie et al., 2019). Thus, following induction of EPHB2 expression, MYCBP2CRISPR and CTRLCRISPR cells were stimulated with pre-clustered Fc control or ephrin-B2 for 15 minutes and scored as collapsed or uncollapsed. While the proportion of collapsed cells for the two cell lines treated with Fc was similar (CTRLCRISPR, 5.4%; MYCBP2CRISPR, 5.7%; p=0.9841), ephrin-B2 treatment resulted in the collapse of 19.0% of CTRLCRISPR cells but only 12.1% of MYCBP2CRISPR cells (n=9 coverslips; Fig. 3B, C; p=0.0017). In addition, time-lapse imaging of MYCBP2CRISPR and CTRLCRISPR cells transiently transfected with an EPHB2-GFP expression plasmid revealed a similar attenuation of ephrin-B2-induced cellular contraction (Fig. 3D, E; p=0.0268). These data argue that MYCBP2 regulates a short-term cellular response evoked by ephrin-B2:EPHB2 signalling.

Revised version: To study the role of MYCBP2 in EphB signalling, we took advantage of an experimental paradigm in which exposure to ephrin-B2 evokes cytoskeletal contraction of HeLa cells expressing EPHB2 (Lahaie et al., 2019). Thus, following induction of EPHB2 expression, MYCBP2CRISPR and CTRLCRISPR cells were stimulated with pre-clustered Fc control or ephrin-B2 for 15 minutes and scored as collapsed or uncollapsed. While the proportion of collapsed cells for the two cell lines treated with Fc was similar (CTRLCRISPR, 5.4%; MYCBP2CRISPR, 5.7%; p=0.9841), ephrin-B2 treatment resulted in the collapse of 19.0% of CTRLCRISPR cells but only 12.1% of MYCBP2CRISPR cells (n=9 coverslips; Fig. 3B, C; p=0.0017). Moreover, when compared to the Fc conditions, MYCBP2CRISPR cells exhibited a less drastic change in collapse rate upon ephrin-B2 treatment (Fig. 3B, C; eB2 vs Fc: CTRLCRISPR, p<0.0001; MYCBP2CRISPR, p=0.0115).

Added in the revised version: In addition, time-lapse imaging of MYCBP2CRISPR and CTRLCRISPR cells transiently transfected with an EPHB2-GFP expression plasmid revealed a similar attenuation of ephrin-B2-induced cellular contraction (Fig. 3D, E; p=0.0268). These data argue that MYCBP2 regulates a short-term cellular response evoked by ephrin-B2:EPHB2 signalling.

Earlier version: To study longer-term cellular responses evoked by ephrin-B2:EPHB2 signalling, we turned to a stripe assay in which the preference of CTRLCRISPR and MYCBP2CRISPR cells for immobilized ephrin-B2 or Fc was measured. To do this, cells of either line were deposited over alternating stripes of ephrin-B2 or Fc, and stripe preference was scored for individual cells after overnight incubation. While only 33.2% of CTRLCRISPR cells resided on ephrin-B2 stripes, this proportion was significantly increased to 47.4% for MYCBP2CRISPR cells, suggesting the loss of MYCBP2 function led to a decreased repulsion from ephrin-B2 stripes (Fig. 3F, G; n=5 and 7 carpets respectively, p=0.0109). Together, our data suggest that MYCBP2 is required for EPHB2-mediated cellular responses in HeLa cells.

Revised version: To study longer-term cellular responses evoked by ephrin-B2:EPHB2 signalling, we turned to a stripe assay in which the preference of CTRLCRISPR and MYCBP2CRISPR cells for immobilized ephrin-B2 or Fc was measured. To do this, cells of either line were deposited over alternating stripes of ephrin-B2 or Fc, and stripe preference was scored for individual cells after overnight incubation. While only 33.2% of CTRLCRISPR cells resided on ephrin-B2 stripes, this proportion was significantly increased to 47.4% for MYCBP2CRISPR cells, suggesting the loss of MYCBP2 function led to a decreased repulsion from ephrin-B2 stripes (Fig. 3F, G; n=5 and 7 carpets respectively, p=0.0109). When cells were plated on Fc:Fc stripes, CTRLCRISPR and MYCBP2CRISPR cells exhibited no preference over cy3-conjugated Fc stripes (49.18% vs 51.88%, p=0.5386, images not shown). Native HeLa cells only respond to ephrin-B2 once they are made to express EphB2. Thus, our data suggest that MYCBP2 is required for EPHB2-mediated cellular responses in HeLa cells.

Loss of MYCBP2 decreases cellular levels of EPHB2 protein

Earlier version: The association of EPHB2 with MYCBP2 and its substrate recognition protein FBXO45 suggests that the MYCBP2 ubiquitin ligase complex could target EPHB2 for degradation, a mechanism frequently deployed to terminate transmembrane receptor signalling (Foot et al., 2017). However, the results of our cellular assays contradicted this model, and rather suggested that loss of MYCBP2 function decreased EPHB2 signalling. To further evaluate these two scenarios, we first compared tetracycline-induced levels of EPHB2 protein in HeLa MYCBP2CRISPR cells and CTRLCRISPR cells. We found that impairing MYCBP2 reduced EPHB2 protein levels (Fig. 4A). To extend our data, we transfected EPHB2-FLAG into both CTRLCRISPR and MYCBP2CRISPR cells and examined EPHB2-FLAG levels after two days. As shown in Fig. 4B, C, levels of EPHB2-FLAG were significantly lowered by 26.1% in HeLa MYCBP2CRISPR cells compared to CTRLCRISPR cells (p=0.0046). We further investigated whether MYCBP2 affects EPHB2 protein turnover when cycloheximide is added to prevent new protein synthesis. EPHB2-FLAG expression was induced by tetracycline for 12 hours followed by protein synthesis inhibition with cycloheximide. We found that EPHB2 half-life was reduced in HeLa cells lacking MYCBP2 compared to control (Fig. 4D, E; 8h treatment, p=0.0474).

Revised version: The association of EPHB2 with MYCBP2 and its substrate recognition protein FBXO45 suggests that the MYCBP2 ubiquitin ligase complex could target EPHB2 for degradation, a mechanism frequently deployed to terminate transmembrane receptor signalling (Foot et al., 2017). However, the results of our cellular assays contradicted this model, and rather suggested that loss of MYCBP2 function decreased EPHB2 signalling. To further evaluate these two scenarios, we first compared EPHB2 protein levels in HeLa MYCBP2CRISPR cells and CTRLCRISPR cells by using tetracycline to induce the EPHB2 overexpression. We found that MYCBP2 loss reduced EPHB2 protein levels (Fig. 4A). To further confirm this, instead of inducing EPHB2 overexpression with tetracycline, we transfected EPHB2-FLAG plasmid into both CTRLCRISPR and MYCBP2CRISPR cells and examined EPHB2-FLAG levels after two days. As shown in Fig. 4B, C, levels of EPHB2-FLAG were significantly lowered by 26.1% in HeLa MYCBP2CRISPR cells compared to CTRLCRISPR cells (p=0.0046). We further investigated whether MYCBP2 affects EPHB2 protein turnover when cycloheximide is added to prevent new protein synthesis. EPHB2-FLAG expression was induced by tetracycline for 12 hours followed by protein synthesis inhibition with cycloheximide. We found that EPHB2 half-life was reduced in HeLa cells lacking MYCBP2 compared to control (Fig. 4D, E; 8h treatment, p=0.0474).

MYCBP2 loss-of-function increases EPHB2 protein turnover in HeLa cells.

A) Induced EPHB2-FLAG expression is reduced in MYCBP2CRISPR HeLa cells. B) Western blotting for transfected EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells. C) Quantification of transfected EPHB2-FLAG levels (p=0.0046, one-sample t-test). D) Representative western blot of EPHB2-FLAG in CTRLCRISPR and MYCBP2CRISPR cells treated with DMSO or cycloheximide for 4h and 8h. E) Quantification of EPHB2-FLAG turnover with cycloheximide (CTRLCRISPR vs. MYCBP2CRISPR at 8h eB2-Fc stimulation, p= 0.0474, two-way ANOVA followed by Tukey’s multiple comparison test). F) Western blot showing EPHB2-FLAG degradation when cells are challenged with ephrin-B2 (1µg/ml) for different periods of time. G) Quantification of ephrin-B2-evoked EPHB2 degradation (ns, not significant; two-way ANOVA followed by Tukey’s multiple comparison test). H) Western blot of EPHB2 ubiquitination in CTRLCRISPR and MYCBP2CRISPR cells. I) Quantification of ubiquitinated EPHB2. CTRLCRISPR cells stimulated with Fc vs. eB2-Fc, p= 0.1349 (One sample t-test); MYCBP2CRISPR cells stimulated with Fc vs. EB2-Fc, p= 0.0195 (Unpaired two-tailed t-test). J) After tetracycline induction of EPHB2-FLAG expression for 16hrs, CTRLCRISPR and MYCBP2CRISPR HeLa cells were treated with DMSO (1:500) or inhibitors of the proteasome (MG132 50µM) or lysosome (BafA1 0.2µM; CoQ 50µM) for 6 hours, and EPHB2 levels were analysed by western blotting. K) Quantification of EPHB2 levels following treatment with proteasome or lysosome inhibitors. Statistical significance for the comparison between CTRLCRISPR cells treated with DMSO or inhibitors was determined by one-sample t-test (MG132, p=0.0598; BafA1, p=0.0200; CoQ, p=0.3632), whereas statistical significance for the comparison between MYCBP2CRISPR cells treated with DMSO and individual inhibitors was determined by two-tailed paired t-test (MG132, p=0.8893; BafA1, p=0.0361; CoQ, p=0.0835). L) GFP-EPHB1 and HA-EPHB3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by Western blot. M) Quantification of GFP-EPHB1 and HA-EPHB3 levels (EPHB1, p=0.0588; EPHB3, p=0.4253; one-sample t-test). N) FLAG-EPHA3 transfected into CTRLCRISPR and MYCBP2CRISPR HeLa cells and detected by WB. O) Quantification of FLAG-EPHA3 (p=0.5369, one-sample t-test). Error bars represent SD.

Ephrin ligand treatment eventually results in Eph receptor degradation, a process associated with signalling termination. We therefore asked whether ligand-mediated EPHB2 receptor degradation depends on MYCBP2. We induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells, and exposed them to ephrin-B2 for a different length of time. Western blotting revealed that 4-8 hours stimulation with ephrin-B2 reduced EPHB2 levels in CTRLCRISPR cells, but this effect was more drastic in MYCBP2CRISPR cells (Fig. 4F, G). Taken together, our results are not consistent with MYCBP2 ubiquitinating EPHB2 and causing its degradation. Unexpectedly, our data indicate that MYCBP2 stabilizes EPHB2 in HeLa cells under both naïve and ligand-challenged conditions.

Loss of MYCBP2 enhances ligand-induced EPHB2 receptor ubiquitination

Ubiquitination of receptor tyrosine kinases, including Eph receptors, can herald their degradation via the proteasome and thus termination of signalling (Haglund and Dikic, 2012;Sabet et al., 2015). This model is not supported by our results, which suggest that MYCBP2 is required for EPHB2 protein maintenance. Nonetheless, we investigated whether EPHB2 receptor ubiquitination is altered in HeLa cells depleted of MYCBP2. HeLa CTRLCRISPR and MYCBP2CRISPR cells were transfected with HA-tagged ubiquitin, EPHB2 expression was induced with tetracycline, and cells were treated with ephrin-B2 for 30min. We observed that EPHB2 receptor ubiquitination was not significantly increased in CTRLCRISPR cells after this short-term ligand treatment (Fig. 4H, I CTRLCRISPR, p=0.1349). In contrast, EPHB2 ubiquitination was significantly increased in MYCBP2CRISPR cells (Fig. 4H, I MYCBP2CRISPR, p=0.0195). This effect argues against the concept that EPHB2 is a MYCBP2 ubiquitination substrate, and suggests that in the absence of MYCBP2 degradation of the EPHB2 receptor is enhanced due to increased ubiquitination.

Loss of MYCBP2 promotes EPHB2 degradation through the lysosomal pathway

Earlier version: EPHB2 can be degraded by either a proteasomal or lysosomal pathway depending on the cellular context (Cissé et al., 2011; Fasen et al., 2008; Litterst et al., 2007). Thus, we wanted to shed light on how EPHB2 is degraded and understand why EPHB2 degradation is enhanced by MYCBP2 loss of function. To do so, we induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells and applied the S26 proteasome inhibitor MG132, or the lysosomal inhibitors BafilomycinA1 or Chloroquine. We found that MG132 did not have significant effects on EPHB2 levels in both cell types (Fig. 4J, K). However, we found that BafilomycinaA1 (BafA1) significantly increased EPHB2 protein levels in both HeLa CTRLCRISPR cells and MYCBP2CRISPR cells by 19% and 40%, respectively (Fig. 4J, K). We also observed a trend towards increased EPHB2 levels with Chloroquine (coQ) treatment in CTRLCRISPR (14%) and MYCBP2CRISPR (35%), further suggesting a role for lysosomal degradation (Fig. 4J, K). These observations suggest that the loss of MYCBP2 promotes EPHB2 receptor degradation through the lysosomal pathway, which expands our understanding of how lysosomal degradation of EPHB2 is regulated.

Revised version: EPHB2 can be degraded by either a proteasomal or lysosomal pathway depending on the cellular context (Cissé et al., 2011; Litterst et al., 2007;Fasen et al., 2008). Thus, we wanted to shed light on how EPHB2 is degraded and understand why EPHB2 degradation is enhanced by MYCBP2 loss of function. To do so, we induced EPHB2 expression in CTRLCRISPR and MYCBP2CRISPR cells and applied the S26 proteasome inhibitor MG132, or the lysosomal inhibitors BafilomycinA1 or Chloroquine. We found that MG132 did not have significant effects on EPHB2 levels in both cell types (Fig. 4J, K). However, we found that BafilomycinaA1 (BafA1) significantly increased EPHB2 protein levels in both HeLa CTRLCRISPR cells and MYCBP2CRISPR cells by 19% and 40%, respectively (Fig. 4J, K). We also observed a trend towards increased EPHB2 levels with Chloroquine (coQ) treatment in CTRLCRISPR (14%) and MYCBP2CRISPR (35%), further suggesting a role for lysosomal degradation (Fig. 4J, K). These observations suggest that the loss of MYCBP2 promotes EPHB2 receptor degradation through the lysosomal pathway, which expands our understanding of how lysosomal degradation of EPHB2 is regulated.

Loss of MYCBP2 does not decrease EPHB1, EPHB3, or EPHA3 receptor levels

The above experiments raise the question of whether MYCBP2 is a general regulator of Eph receptor stability. Since EPHA3 does not form a complex with MYCBP2, and EphA receptor levels are controlled by the proteasomal pathway (Sharfe et al., 2003;Walker-Daniels et al., 2002), we hypothesized that MYCBP2 might regulate the levels of the entire EphB receptor class. To test this, we co-transfected plasmids encoding GFP-tagged EPHB1 and HA-tagged EPHB3 into HeLa CTRLCRISPR and MYCBP2CRISPR cells. Compared to CTRLCRISPR cells, EPHB1 and EPHB3 levels were not significantly reduced in MYCBP2CRISPR cells (Fig. 4L, M, EPHB1, p=0.0588; EPHB3, p=0.4253). Likewise, FLAG-EPHA3 levels were similar in CTRLCRISPR and MYCBP2CRISPR cells (Fig. 4N, O, p=0.5369). Taken together, these data suggest that MYCBP2 specifically affects EPHB2 receptor stability.

Loss of MYCBP2 attenuates the magnitude of EPHB2 cellular signalling

EPHB2 receptor activation evokes signal transduction events such as tyrosine phosphorylation of EPHB2, activation of EPHB2 tyrosine kinase function and phosphorylation of the ERK1/2 downstream effector (Poliakov et al., 2008). We therefore measured EPHB2 and ERK1/2 phosphorylation in CTRLCRISPR and MYCBP2CRISPR cells for up to 8 hours after ephrin-B2 application (Fig. 5A). P-EPHB2 (pY20) and p-ERK1/2 signals were normalised to GAPDH and ERK1/2, respectively. In HeLa CTRLCRISPR cells, pTyr-EPHB2 response reached a plateau after 1-2 h treatment and remained up 8 hours post-stimulation (Fig. 5A-C). On the contrary, ephrin-B2-evoked phosphorylation of EPHB2 was reduced in MYCBP2CRISPR cells with quantitative results showing a significant reduction by 8 hours of treatment (Fig. 5A, C; 8h p=0.0331). We were also able to detect significantly lower p-ERK1/2 levels in MYCBP2CRISPR cells relative to CTRLCRISPR cells, although activation of ERK1/2 by ephrin-B2 was variable (Fig. 5D. 4h, p=0.0494; 8h, p=0.0078; n=6). We again noted significantly enhanced EPHB2 degradation in MYCBP2CRISPR cells (Fig. 5E; 8h, p=0.0437). This decrease was also observed when CTRLCRISPR and MYCBP2CRISPR cells were treated with ephrin-B1 (data not shown).

MYCBP2 depletion impairs EPHB2 phosphorylation and ERK1/2 activation in HeLa cells.

A) Representative western blot for pERK1/2 and pTyr-EPHB2 detected in CTRLCRISPR and MYCBP2CRISPR cells treated with ephrin-B2 (eB2-Fc) for different periods (n=6). Membranes were striped and reblotted with anti-ERK1/2, anti-EPHB2, anti-GAPDH and anti-MYCBP2 antibodies as controls. B) Quantification of EPHB2 tyrosine phosphorylation in CTRLCRISPR cells evoked by ephrin-B2 treatment (15 min, p=0.1363; 30 min, p=0.2056; 1h, p=0.0342; 2h, p=0.0234; 4h, p=0.0068; 8h, p= 0.0231; one-sample t-test). C) Quantification of EPHB2 tyrosine phosphorylation in CTRLCRISPR and MYCBP2CRISPR HeLa cells (unstimulated, p=0.5589, one-sample t-test; stimulated for 15 min, p=0.7463; 30 min, p=0.5520; 1h, p=0.1920; 2h, p=0.2009; 4h, p=0.1550; 8h, p=0.0331; two-tailed unpaired t-test). D) Quantification of pERK1/2 in CTRLCRISPR and MYCBP2CRISPR HeLa cells (0 min, p=0.0168, one-sample t-test; 15 min, p=0.6695; 30 min, p= 0.6649; 1h, p= 0.1776; 2h, p= 0.1479; 4h, p= 0.0494; 8h, p= 0.0078; two-tailed unpaired t-test). E) Quantification of EPHB2 in CTRLCRISPR and MYCBP2CRISPR HeLa (0 min, p=0.4604, one-sample t-test; 15 min, p=0.2222; 30 min, p=0.1376; 1h, p=0.0651; 2h, p=0.0736; 4h, p=0.2451; 8h, p=0.0437, two-tailed unpaired t-test). F) Quantification of ephrin-B2-evoked EPHB2 tyrosine phosphorylation levels relative to total EPHB2 protein levels (0 min, p=0.7058, one-sample t-test; 15 min, p=0.2464; 30 min, p=0.7835; 1h, p=0.9164; 2h, p=0.7196; 4h, p=0.8625; 8h, p=0.5750, two-tailed unpaired t-test). G) Quantification of ephrin-B2-evoked pERK1/2 relative to EPHB2 total protein levels (0 min, p=0.3308, one-sample t-test; 15 min, p=0.3856; 30 min, p=0.0624; 1h, p=0.2683; 2h, p=0.1284; 4h, p=0.7998; 8h, p=0.7790, two-tailed unpaired t-test. Error bars represent SD.

Next, we asked whether the decrease in ligand-evoked EPHB1/2 and ERK1/2 activation in MYCBP2CRISPR cells reflects a requirement for MYCBP2 in the EPHB2 signalling cascade per se, or whether it is explained by decreased EPHB2 protein levels caused by MYCBP2 loss. We thus normalized p-EPHB2 and p-ERK1/2 signal to the levels of EPHB2 protein at all time points, which revealed that kinetics and magnitude of EPHB2 and ERK1/2 phosphorylation are similar between CTRLCRISPR and MYCBP2CRISPR cells (Fig. 5F, G). Although these data argue against a direct role of MYCBP2 in the early events of EPHB2 signalling, they nevertheless indicate that MYCBP2 loss results in marked attenuation of EPHB2 activation and its downstream pERK1/2 signalling in line with decreased cellular responses to ephrin-B2.

Exogenous Fbxo45 binding domain of MYCBP2 disrupts the EPHB2-MYCBP2 interaction

Previous studies showed that FBXO45 binds to the FBD1 domain of MYCBP2, and exogenous FBD1 overexpression can disrupt the FBXO45-MYCBP2 association (Sharma et al., 2014). We thus tested whether FBD1 overexpression can interfere with the formation of the EPHB2-MYCBP2 complex (Fig. 6A). Indeed, the expression of GFP-FBD1 wild-type (WT) in HEK cells expressing EPHB2-FLAG and MYC-FBXO45 reduced binding between EPHB2 and MYCBP2 (Fig. 6B). This effect was not observed with a GFP-FBD1 mutant (mut) fragment that harbours three point mutations that inhibit binding to FBXO45. Reduction of the EPHB2-MYCBP2 interaction was also observed in cells that were not overexpressing FBXO45 (Fig. 6C). Thus, exogenous FBD1 can specifically disrupt the EPHB2-MYCBP2 association. This suggests that FBXO45 binding to MYCBP2 may facilitate for formation of the MYCBP2-EPHB2 complex.

Exogenous FBD1 fragment of MYCBP2 disrupts EPHB2-MYCBP2 binding and impairs EPHB2 function in HeLa cells.

A) Schematic illustrating competition of exogenous MYCBP2-FBD1 fragment that disrupts MYCBP2-FBXO45 binding and leads to MYCBP2 reduction in EPHB2 complexes. B) Exogenous FBD1 WT overexpression leads to reduced EPHB2-MYCBP2 binding in HEK293 cells despite co-expression of FBXO45. C) FBD1 overexpression also disrupts EPHB2-MYCBP2 binding in the absence of FBXO45 overexpression. D) Representative images of ephrin-B2 stripe assays using HeLa cells expressing GFP-FBD1 mut or GFP-FBD1 WT. E) Quantification of cells present on eB2 stripes (p=0.0107, two-tailed unpaired t-test). Error bars represent SD.

FBD1 expression impairs EPHB2-mediated neuronal responses

Earlier version: To determine whether the EPHB2-MYCBP2 interaction is required for EPHB2 function, we disrupted binding via FBD1 ectopic expression in cell lines or primary neurons, and studied cellular responses to ephrin-B treatment. We introduced GFP-FBD1 mut or GFP-FBD1 WT into HeLa cells with inducible EPHB2 expression and cultured these cells on ephrin-B2 and Fc stripes. Following overnight culture, only 19.2% of cells expressing GFP-FBD1 mut resided on ephrin-B2 stripes. In contrast, 28.6% of cells expressing GFP-FBD1 WT were found on ephrin-B2 stripes indicating that FBD1 expression can dampen ephrin-B2:EPHB2 mediated cell repulsion (Fig. 6D, E; p=0.0107).

Revised version: To determine whether the EPHB2-MYCBP2 interaction is required for EPHB2 function, we disrupted binding via FBD1 ectopic expression in cell lines or primary neurons, and studied cellular responses to ephrin-B treatment. We introduced GFP-FBD1 mut or GFP-FBD1 WT into HeLa cells with inducible EPHB2 expression and cultured these cells on ephrin-B2 and Fc stripes. Following overnight culture, only 19.2% of cells expressing GFP-FBD1 mut resided on ephrin-B2 stripes. In contrast, 28.6% of cells expressing GFP-FBD1 WT were found on ephrin-B2 stripes indicating that FBD1 expression can dampen ephrin-B2:EPHB2 mediated cell repulsion (Fig. 6D, E; p=0.0107). In contrast, cells expressing GFP-FBD1 mut or GFP-FBD1 WT displayed no preference for either one of Fc:Fc control stripes (49.63% vs 49.85%, p=0.9560, images not shown).

Since EPHB2 is expressed in the embryonic chicken spinal cord, we performed in ovo electroporation to introduce plasmids encoding GFP-FBD1 mut or GFP-FBD1 WT into the spinal cords of 2.5-day old embryonic chickens (Hamburger-Hamilton; (HH) Stage 15-17; (Hamburger and Hamilton, 1951;Luria et al., 2008)). Two days later spinal cords were dissected, divided into explants, cultured overnight on alternating stripes containing ephrin-B2 or Fc, and axonal GFP signal present over ephrin-B2 versus Fc stripes was determined. When explants expressed GFP-FBD1 mut (negative control), we observed 26.5% of axonal GFP signal resided on ephrin-B2 stripes (Fig. 7A, B). In contrast, we found a significant increase of 32.2% of neurites expressing FBD1 WT on ephrin-B2 stripes (Fig. 7A, B; p=0.0410). Thus, FBD1 expression impairs long-term repulsive responses to ephrinB2-EPHB2 signalling in spinal explants.

Exogenous FBD1 overexpression impairs EPH receptor functions in chick spinal cord explants and mouse hippocampal neurons.

Earlier version: A) Representative images of ephrin-B2 stripe assays with chick embryonic spinal cord explants overexpressing GFP-FBD1 mut (negative control) or GFP-FBD1 WT. B) Quantification of GFP-positive branches present on ephrin-B2 stripes (GFP-FBD1 mut vs. GFP-FBD1 WT, p=0.0410, two-tailed unpaired t-test). C) Representative images of DIV2 mouse hippocampal neurons overexpressing GFP-FBD1 mut or GFP-FBD1 WT and challenged with Fc control, ephrin-B1 (eB1-Fc) or ephrin-B2 (eB2-Fc). D) Quantification of growth cone collapse rate for hippocampal neurites. GFP-FBD1 mut: Fc vs. eB1-Fc, p=<0.0001; Fc vs. eB2-Fc, p=0.0008; GFP-FBD1 WT: Fc vs. eB1-Fc p=0.4357; Fc vs eB2-Fc p=0.0065. Two-way ANOVA followed by Tukey’s multiple comparisons test. Error bars represent SD.

Revised version: A) Representative images of ephrin-B2 stripe assays with chick embryonic spinal cord explants overexpressing GFP-FBD1 mut (negative control) or GFP-FBD1 WT. Images with inverted GFP signal in dark pixels on Fc (white) / eB2 (pink) stripes are placed beside the original images. B) Quantification of GFP-positive branches present on ephrin-B2 stripes (GFP-FBD1 mut vs. GFP-FBD1 WT, p=0.0410, two-tailed unpaired t-test). C) Representative images of DIV2 mouse hippocampal neurons overexpressing GFP-FBD1 mut or GFP-FBD1 WT and challenged with Fc control, ephrin-B1 (eB1-Fc) or ephrin-B2 (eB2-Fc). D) Quantification of growth cone collapse rate for hippocampal neurites. GFP-FBD1 mut: Fc vs. eB1-Fc, p=<0.0001; Fc vs. eB2-Fc, p=0.0008; GFP-FBD1 WT: Fc vs. eB1-Fc p=0.4357; Fc vs eB2-Fc p=0.0065. Two-way ANOVA followed by Tukey’s multiple comparisons test. Error bars represent SD.

To study short-term neuronal responses to ephrin-Bs, we turned to mouse hippocampal neurons and a growth cone collapse assay (Srivastava et al., 2013). Here, we electroporated GFP-FBD1 mut or GFP-FBD1 WT expression plasmids into dissociated hippocampal neurons and treated them with pre-clustered Fc, ephrin-B1 or ephrin-B2 for one hour (Fig. 7C). Neurons expressing GFP-FBD1 mut showed 33.9% growth cone collapse with Fc treatment, while Ephrin-B1 and Ephrin-B2 treatment elicited significant increases to 57.9% and 53.3% collapse, respectively (Fig. 7C, D). Importantly, both ephrin-B ligands significantly increased the frequency of growth cone collapse with respect to Fc treatment when FBD1 mut was expressed (Fig. 7D; p<0.0001, and p=0.0008). In contrast, this effect was either significantly reduced or less evident in neurons expressing GFP-FBD1 WT: with only 48.1% and 57.5% growth cones being collapsed by ephrin-B1 and ephrin-B2, respectively, compared to 41.7% being collapsed by Fc treatment (Fig. 7D; p=0.4357 and p=0.0065, respectively). Together, these data indicate that impairing MYCBP2 function via FBD1 expression disrupts ephrin:B-EPHB2 signalling in axonal guidance.

Removed in the revised version: C. elegans ephrin receptor VAB-1 interacts genetically with the RPM-1/MYCBP2 signalling network

Removed in the revised version: Next, we sought to test genetic interactions between an ephrin receptor and the MYCBP2 pathway using an in vivo animal model. To do so, we turned to C. elegans where we evaluated the sole Eph family receptor in this organism, VAB-1, and its genetic relationship with several genes in the MYCBP2 signalling network. C. elegans has a single MYCBP2 ortholog called RPM-1, and previous studies have shown that RPM-1/MYCBP2 is required to terminate axon outgrowth (Grill et al., 2016). Previous work showed that axon termination defects in rpm-1 mutants are due to impaired growth cone collapse (Borgen et al., 2017). Axon termination is also regulated by several downstream players in the RPM-1 signalling network including FSN-1/FBXO45 and GLO-4, a Rab GEF orthologue of human SERGEF (Grill et al., 2007).

Removed in the revised version: To evaluate genetic interactions between the VAB-1 Eph receptor and RPM-1/MYCBP2 signalling, we used the PLM mechanosensory neurons of C. elegans. There are two PLM neurons with one on either side of the body. Each PLM neuron extends a primary axon that grows anteriorly and terminates at a precise location before the cell body of another mechanosensory neuron called ALM (Fig. 8A). Consistent with previous studies (Mohamed and Chin-Sang, 2006), vab-1 loss-of-function (lf) mutants presented moderate axon termination defects where the PLM axon overextends beyond the ALM cell body (Fig. 8A). Quantitation indicated that this phenotype occurred at a low, but significant frequency (Fig. 8B). We then examined the genetic relationship between vab-1 and the Rab GEF glo-4. We observed both moderate overextension defects and more severe defects in which the axon hooked ventrally in vab-1; glo-4 double mutants (Fig. 8A). The frequency of both types of axon termination defects were significantly enhanced in vab-1; glo-4 double mutants, and overextension defects were significantly rescued by transgenic expression of VAB-1 (Fig. 8A, B). We subsequently tested the relationship between the Fbox protein fsn-1 and vab-1, and observed enhanced defects in vab-1; fsn-1 double mutants (Fig. 8C). Finally, we evaluated vab-1; rpm-1 double mutants and did not observe a significant change in phenotypic severity or penetrance compared to rpm-1 single mutants (Fig. 8D).

Earlier version: Vab-1 ephrin receptor interacts genetically with MYCBP2/rpm-1 in C. elegans.

Revised version: C. elegans ephrin receptor VAB-1 interacts genetically with known RPM-1/MYCBP2 binding proteins

Added in the revised version: Next, we sought to test genetic interactions between an ephrin receptor and the MYCBP2 signalling network using an in vivo animal model. To do so, we turned to C. elegans which has a sole Eph family receptor called VAB-1(George et al., 1998), and a single MYCBP2 ortholog called RPM-1(Grill et al., 2016). Previous studies have shown that RPM-1/MYCBP2 is required to terminate axon outgrowth with failed termination defects in rpm-1 mutants arising due to impaired growth cone collapse(Schaefer et al., 2000;Borgen et al., 2017). Prior studies have also demonstrated that RPM-1 functions via numerous downstream pathways, and that RPM-1 binding proteins display genetic enhancer interactions (Grill et al., 2007;Grill et al., 2012;Tulgren et al., 2014;Baker et al., 2014). Given our cell-based biochemical results showing that MYCBP2 binds EphB2, we opted to evaluate two known RPM-1 binding proteins for genetic interactions with the VAB-1 Ephrin receptor. 1) The Rab GEF GLO-4 is an RPM-1 binding protein that is orthologous to mammalian SERGEF. GLO-4 functions via the GLO-1/RAB32 small GTPase and is not involved in RPM-1 ubiquitin ligase activity (Grill et al., 2007;Sharma et al., 2014)(Fig 8A). 2) We also tested the C. elegans FBXO45 ortholog, FSN-1, which is the F-box protein of the RPM-1 ubiquitin ligase complex (Liao et al., 2004;Sharma et al., 2014;Baker and Grill, 2017;Desbois et al., 2022) (Fig 8A).

Added in the revised version: C. elegans VAB-1 ephrin receptor interacts genetically with two known RPM-1/MYCBP2 binding proteins FSN-1/FBXO45 and GLO-4/SERGEF.

Earlier version: A) Schematic of axon morphology and axon termination site for PLM mechanosensory neurons. Shown are representative images of PLM axon termination observed for indicated genotypes. Moderate overextension defects (arrowhead) were observed in vab-1 and glo-4 single mutants. Both overextension defects (arrowhead) and more severe hook defects (arrow) were observed in vab-1; glo-4 double mutants. B) Quantitation of axon termination defects for indicated genotypes. vab-1; glo-4 double mutants show enhanced frequency of both severe hook (black) and moderate overextension (grey) axon termination defects. Overextension defects are significantly reduced by transgenic expression of VAB-1. C) Quantitation of axon termination defects for indicated genotypes. vab-1; fsn-1 double mutants show enhanced termination defects. D) Quantitation of axon termination defect for indicated genotypes. Axon termination defects are not enhanced nor suppressed in vab-1; rpm-1 double mutants compared to rpm-1 single mutants. n is defined as a single count of 20-30 animals. Means are shown from 8 to 10 counts (20-30 animals per count) for each genotype, and error bars represent SEM. Significance determined using Student’s t-test with Bonferroni correction. ** p < 0.01; *** p < 0.001; n.s, not significant. Scale bar is 20 μm.

Revised version: A) Schematic shows the known RPM-1/MYCBP2 binding proteins GLO-4/SERGEF and FSN-1/FBXO45. GLO-4 functions independent of RPM-1 ubiquitin ligase. FSN-1 is the F-box protein that forms a ubiquitin ligase complex with RPM-1. Adapted from Grill et al., 2016. B) Schematic of axon morphology and axon termination site for PLM mechanosensory neurons. Shown are representative images of failed axon termination defects observed for PLM neurons of indicated genotypes. Examples of moderate overextension defects (arrowhead) observed in vab-1/EphR and glo-4/SERGEF single mutants. Example of severe hook defects (arrow) observed in vab-1; glo-4 double mutants. C) Quantitation of axon termination defects for indicated genotypes. vab-1; glo-4 double mutants show enhanced frequency of both severe hook (black) and moderate overextension (grey) failed termination defects. Overextension defects are significantly reduced by transgenic expression of VAB-1. D) Quantitation of axon termination defects for indicated genotypes. vab-1; fsn-1 double mutants show enhanced termination defects. E) Quantitation of axon termination defect for indicated genotypes. Axon termination defects are not suppressed in vab-1; rpm-1 double mutants compared to rpm-1 single mutants. n is defined as a single count of 20-30 animals. Means are shown from 8 to 10 counts (20-30 animals per count) for each genotype, and error bars represent SEM. Significance determined using Student’s t-test with Bonferroni correction. ** p < 0.01; *** p < 0.001; n.s, not significant. Scale bar is 20 μm.

Added in the revised version: To evaluate genetic interactions between the VAB-1 Eph receptor and RPM-1/MYCBP2 binding proteins, we used the PLM mechanosensory neurons of C. elegans. There are two PLM neurons with one on either side of the body. Each PLM neuron extends a primary axon that grows anteriorly and terminates at a precise location before the cell body of another mechanosensory neuron called ALM (Fig. 8B). Consistent with previous studies (Mohamed and Chin-Sang, 2006), vab-1/EphR loss-of-function (lf) mutants presented moderate failed termination defects where the PLM axon overextends beyond the ALM cell body (Fig. 8B).

Added in the revised version: Quantitation indicated that this phenotype occurred at a low, but significant frequency (Fig. 8C). We then examined the genetic relationship between vab-1 and the RPM-1 binding protein glo-4/SERGEF. In vab-1; glo-4 double mutants, we observed both moderate overextension defects and more severe failed termination defects where the axon hooked ventrally (Fig. 8B). The frequency of both types of axon termination defects were significantly enhanced in vab-1; glo-4 double mutants (Fig. 8B, C). Overextension defects were significantly rescued by transgenic expression of VAB-1 in vab-1; glo-4 double mutants (Fig. 8C). We also tested the relationship between the F-box protein fsn-1/FBXO45 and vab-1, and observed enhanced defects in vab-1; fsn-1 double mutants (Fig. 8D). Finally, we evaluated vab-1; rpm-1 double mutants and did not observe a significant change in phenotypic severity or penetrance compared to rpm-1 single mutants (Fig. 8E).

Earlier version: Collectively, these genetic experiments support several conclusions. 1) Enhanced defects in vab-1; glo-4 and vab-1; fsn-1 double mutants indicate that the VAB-1 Ephrin receptor functions in parallel to known players in the RPM-1 signalling network to facilitate axon termination. 2) The absence of enhanced defects in vab-1; rpm-1 double mutants indicates that VAB-1 functions in the same genetic pathway as RPM-1. 3) Lack of suppression in vab-1; fsn-1 and vab-1; rpm-1 double mutants indicates that VAB-1 is neither ubiquitinated nor inhibited/degraded by the RPM-1/FSN-1 ubiquitin ligase complex. Together, these data from the C. elegans model are consistent with our proteomic, biochemical and cellular observations arguing that MYCBP2 binds to EPHB2 to protect it from degradation.

Revised version: Collectively, these genetic experiments support two conclusions. 1) Enhanced defects in vab-1; glo-4 and vab-1; fsn-1 double mutants indicate that the VAB-1 Ephrin receptor functions in parallel to known RPM-1 binding proteins to facilitate axon termination. These results are consistent with the VAB-1 Ephrin receptor functioning within the RPM-1 signalling network. 2) Lack of suppression in vab-1; fsn-1 and vab-1; rpm-1 double mutants indicates that VAB-1 is not inhibited or degraded by the RPM-1/FSN-1 ubiquitin ligase complex. Thus, our findings in C. elegans are consistent with our biochemical studies in mammalian cells that demonstrate MYCBP2 stabilizes EPHB2 rather than degrading it. We note that because C. elegans contains only a single Ephrin receptor, our findings with C. elegans do not necessarily pertain specifically to EPHB2. Nonetheless, these results add further evidence indicating there are substantial genetic interactions between known MYCBP2 binding proteins and Eph receptors.

Discussion

Our MS-based proteomics efforts to identify EPHB2 interacting proteins yielded MYCBP2, a signalling hub and ubiquitin ligase that is functionally linked to many of the cellular processes also mediated by EPHB2, including cellular growth, proliferation, synapse formation, and axon development. Our experiments argue against EPHB2 being a MYCBP2 ubiquitination substrate. Instead, we envisage a model where MYCBP2 controls EPHB2 signalling indirectly by preventing its lysosomal degradation and maintaining EPHB2 protein levels sufficiently high to mediate efficient cellular and axonal growth cone repulsion from ephrin-B ligands. The interaction between MYCBP2 and EPHB2 may allow the coupling EPHB2 to fundamental cellular processes that control growth, proliferation, and survival. Here, we discuss the molecular logic of MYCBP2 and EPHB2 association in the context of Eph receptor signalling, its potential implications for neural development and diversification of EphB2 signalling.

Functional significance of MYCBP2-EPHB2 complex formation

Earlier version: Our biochemical experiments validate the formation of a MYCBP2-EPHB2 complex and suggest that this association is decreased following ligand application. Because PHR proteins like MYCBP2 are large signalling hubs, MYCBP2 association with EPHB2 might sterically hinder the formation of EPHB2 multimers and clusters necessary for signalling. Thus, ligand-induced dissociation of the MYCBP2-EPHB2 complex could be a prelude to signalling. Our findings indicate that MYCBP2 association with EPHB2 is enhanced by FBXO45, a subunit of the MYCBP2/FBXO45 complex that mediates ubiquitination substrate binding. Previous work showed that MYCBP2 functions to polyubiquitinate specific protein substrates, targeting them for degradation and inhibition (Crawley et al., 2019; Desbois et al., 2022; Han et al., 2012; Nakata et al., 2005). We considered the possibility that the ubiquitin ligase activity of MYCBP2 is important for the termination of EPHB2 signalling, but several lines of evidence argue against this idea: (1) the application of ephrin-B2 ligand results in the dissociation of the MYCBP2-EPHB2 complex, (2) loss of MYCBP2 function results in decreased cellular levels of EPHB2 protein, and (3) impairing MYCBP2 increases ligand-stimulated EPHB2 ubiquitination. In addition, our quantification of EPHB2 signalling suggested that decreased EPHB2 and ERK phosphorylation seen in MYCBP2-deficient cells can be accounted for by the decrease in EPHB2 receptor levels. Thus, a more plausible model that is consistent with our results is that MYCBP2 association with EPHB2 protects EPHB2 from turnover by lysosome-mediated degradation and prevents EPHB2 ubiquitination by the action of unidentified ubiquitin ligases. Interestingly, studies in C. elegans have shown that RPM-1/MYCBP2 regulates lysosome biogenesis (Grill et al., 2007). While this established a link between MYCBP2 signalling and the endo-lysosomal degradation system, our results now indicate that MYCBP2 can also influence turnover of EPHB2 via lysosomal degradation. Consistent with these findings, we observed that the Eph receptor VAB-1 functions in the RPM-1 signalling network and in a parallel genetic pathway with GLO-4, the RPM-1 binding protein that mediates effects on lysosome biogenesis.

Revised version: Our biochemical experiments validate the formation of a MYCBP2-EPHB2 complex and suggest that this association is decreased following ligand application. Because PHR proteins like MYCBP2 are large signalling hubs, MYCBP2 association with EPHB2 might sterically hinder the formation of EPHB2 multimers and clusters necessary for signalling. Thus, ligand-induced dissociation of the MYCBP2-EPHB2 complex could be a prelude to signalling. Our findings indicate that MYCBP2 association with EPHB2 is enhanced by FBXO45, a subunit of the MYCBP2/FBXO45 complex that mediates ubiquitination substrate binding. Previous work showed that MYCBP2 functions to polyubiquitinate specific protein substrates, targeting them for degradation and inhibition (Crawley et al., 2019;Han et al., 2012;Nakata et al., 2005;Desbois et al., 2022). We considered the possibility that the ubiquitin ligase activity of MYCBP2 is important for the termination of EPHB2 signalling, but several lines of evidence argue against this idea: (1) the application of ephrin-B2 ligand results in the dissociation of the MYCBP2-EPHB2 complex (Walker-Daniels et al.), (2) loss of MYCBP2 function results in decreased cellular levels of EPHB2 protein, and (3) impairing MYCBP2 increases ligand-stimulated EPHB2 ubiquitination. In addition, our quantification of EPHB2 signalling suggested that decreased EPHB2 and ERK phosphorylation seen in MYCBP2-deficient cells can be accounted for by the decrease in EPHB2 receptor levels. Thus, a more plausible model that is consistent with our results is that MYCBP2 association with EPHB2 protects EPHB2 from turnover by lysosome-mediated degradation and prevents EPHB2 ubiquitination by the action of unidentified ubiquitin ligases. Interestingly, studies in C. elegans have shown that RPM-1/MYCBP2 regulates lysosome biogenesis via the GLO-4/GLO-1 pathway (Grill et al., 2007). While this established a link between MYCBP2 signalling and the endo-lysosomal degradation system, our results now indicate that MYCBP2 can also influence turnover of EPHB2 via lysosomal degradation. Consistent with these findings, we observed that the Eph receptor VAB-1 functions in parallel genetic pathways with the known RPM-1 binding protein GLO-4/SERGEF that regulates lysosome biogenesis.

The “protective” effect of MYCBP2 vis-à-vis EPHB2 ubiquitination and lysosomal degradation might be secondary to effects on other aspects of EPHB2 signalling that we have not explored experimentally. For example, since MYCBP2 is a signalling hub with multiple substrates and binding proteins, it could bring EphB2 into close proximity to components of the MYCBP2 signalling network integrating cell-cell communication via ephrin:Eph signals with MYCBP2 intracellular signalling to influence fundamental cellular processes. One example of this is suggested by a recent study demonstrating a link between EPHB2 and mTOR-mediated cell growth signalling pathways (Gagne et al., 2021). One potential mechanism could involve Ephrin binding-induced dissociation of the EPHB2-MYCBP2 interaction, allowing MYCBP2 ubiquitination of the TSC complex that regulates mTOR function (Han et al., 2012).

The formation of a MYCBP2-EPHB2 complex that includes other signalling receptors could also explain the curious result that the extracellular domain of EPHB2 is critical for MYCBP2 association. The finding would be consistent with the extracellular domain of EPHB2 interacting with the extracellular domains of other receptors, whose intracellular domains are linked more directly with MYCBP2 and FBXO45. An alternative explanation may be a non-classical extracellular MYCBP2-EPHB2 interaction similar to the one proposed for the extracellular domain of N-Cadherin and FBXO45 (Na et al., 2020).

MYCBP2 and EPHB2 functions in the developing nervous system

Earlier version: Our proteomic, biochemical, and genetic experiments indicate that MYCBP2 and EPHB2 function in the same pathway. This is also supported by the striking similarity of MYCBP2 and EPHB2 loss of function phenotypes in the mouse nervous system. In the context of developing neuronal connections, EphB2 and Mycbp2/Phr1 mouse mutants display similar axon guidance phenotypes such as abnormal limb nerve trajectories by motor axons, defective growth cone crossing of the midline at the level of the optic chiasm and decreased connectivity between the two cortical hemispheres through the corpus callosum (D’Souza et al., 2005; Henkemeyer et al., 1996; Lewcock et al., 2007; Luria et al., 2008; Williams et al., 2003). Both proteins also function in synaptic development and their loss of function leads to decreased numbers of synapses and altered synapse morphology (Bloom et al., 2007; Dalva et al., 2000; Wan et al., 2000; Zhen et al., 2000). Furthermore, outside the nervous system, both are involved in a variety of cancers with recent evidence linking them to esophageal adenocarcinoma and c-MYC-dependent control of cell proliferation (Genander et al., 2009; Han et al., 2012; Venkitachalam et al., 2022). Our proteomic, biochemical, and cellular experiments together with genetic experiments in C. elegans now provide a new framework in which to consider phenotypic and disease links between EPHB2 and MYCBP2.

Revised version: Our proteomic, biochemical, and genetic experiments indicate that MYCBP2 and EPHB2 function in the same pathway. This is also supported by the striking similarity of MYCBP2 and EPHB2 loss of function phenotypes in the mouse nervous system. In the context of developing neuronal connections, EphB2 and Mycbp2/Phr1 mouse mutants display similar axon guidance phenotypes such as abnormal limb nerve trajectories by motor axons, defective growth cone crossing of the midline at the level of the optic chiasm and decreased connectivity between the two cortical hemispheres through the corpus callosum (Lewcock et al., 2007;Luria et al., 2008;Williams et al., 2003;Henkemeyer et al., 1996;D’Souza et al., 2005). Both proteins also function in synaptic development and their loss of function leads to decreased numbers of synapses and altered synapse morphology (Wan et al., 2000;Bloom et al., 2007;Dalva et al., 2000;Zhen et al., 2000). Furthermore, outside the nervous system, both are involved in a variety of cancers with recent evidence linking them to esophageal adenocarcinoma and c-MYC-dependent control of cell proliferation (Venkitachalam et al., 2022;Han et al., 2012;Genander et al., 2009). Our proteomic, biochemical, and cellular experiments together with genetic interaction studies in C. elegans now provide a new framework in which to consider phenotypic and disease links between EPHB2 and MYCBP2.

Importantly, the biomedical relevance of our findings are heightened by a recent study that identified genetic variants in MYCBP2, which cause a neurodevelopmental disorder termed MYCBP2-related Developmental delay with Corpus callosum Defects (MDCD) (Al-Abdi et al., 2022). MDCD features defective neuronal connectivity including a hypoplastic or absent corpus callosum, neurobehavioral deficits including intellectual disability and epilepsy, and abnormal craniofacial development. This constellation of comorbidities in MDCD closely resembles some of the phenotypes observed in mice with deficient EPHB2 signalling. Given our finding that MYCBP2 loss reduces EPHB2 levels and influences Eph receptor effects on axon development, the question of whether EPHB2 expression levels are normal in MYCBP2 patients remains pertinent.

In conclusion, our study has revealed numerous biochemical and genetic links between MYCBP2 and EPHB2. Our findings indicate that the MYCBP2/FBXO45 complex protects EPHB2 from degradation, and these are functionally integrated signalling players with an evolutionarily conserved role in axonal development. Future studies will be needed to address how the EPHB2-MYCBP2 interaction affects nervous system development in mammals in vivo and to identify further regulators of EPHB2 degradation. Additionally, another idea worthy of closer examination in the future is the possibility that MYCBP2 signalling could provide routes through which EPHB2-initiated signals access numerous fundamental cellular functions.